# <center>Computational Physics</center>
---

## Week 2: Numerical Integration

```python
import numpy 
import matplotlib.pyplot as plt
%matplotlib inline
```

Define the function `f`, such that $\textrm{f}(x) \equiv x^{2}\sin(x)$. This is the function that we will be integrating.

<!-- grade_id: integrand, points: 0, grade: false, solution: true -->
```python
import numpy as np
def f(x):
    '''Function equivalent to x^2 sin(x).'''
    # YOUR CODE HERE
    return((x**2)*np.sin(x))
```

Ensure your function works with numpy arrays:

<!-- grade_id: cell-653afd49e7782cc4, points: 1, grade: true, solution: false -->
```python
xs=numpy.arange(0, 1, step=0.1)
assert numpy.isclose(f(xs), 
                     [0., 0.00099833, 0.00794677, 0.02659682, 0.06230693,
                      0.11985638, 0.20327129, 0.31566667, 0.4591079 , 0.6344948 ]).all()
```

Derive the indefinite integral of $f(x)$ nalytically. Call this function $g(x)$ and implement it below. Set the constant of integration such that $g(0)=0$.

<!-- grade_id: integral, points: 0, grade: false, solution: true -->
```python

def g(x):
    '''Analytical integral of f(x).'''
    k=0
    return(-x**2*np.cos(x)+2*x*np.sin(x)+2*np.cos(x)-2)


```

Check your solution with the same numpy array:

<!-- grade_id: cell-648a123ee323b45e, points: 1, grade: true, solution: false -->
```python
assert g(0) == 0.
```

<!-- grade_id: cell-f8eaec03e0540a78, points: 1, grade: true, solution: false -->
```python
assert numpy.isclose(g(xs), 
                     [0., 0.00002497, 0.00039822, 0.00200482, 0.0062869, 
                      0.01519502, 0.03112138, 0.05681646, 0.09529087, 0.1497043 ]).all()
```

Now, using the analytically derived indefinite integral, $g(x)$, define a function which calculates the definite integral of $f(x)$ over the interval $(x_{min},~x_{max})$.

<!-- grade_id: cell-6f8d8cc5f5c7564c, points: 0, grade: false, solution: true -->
```python
def integrate_analytic(xmin, xmax):
    '''Analytical integral of f(x) from xmin to xmax.'''
    # YOUR CODE HERE
    return (g(xmax)-g(xmin))


```

Check your analytic function:

<!-- grade_id: analytical_integral, points: 1, grade: true, solution: false -->
```python
assert numpy.isclose(integrate_analytic(xmin=0, xmax=4), 1.096591)
```

## Numerical implementation

Create a function which calculates the definite integral of the function $f(x)$ over the interval $(x_{min},~x_{max})$ using Simpson's rule with $N$ panels.

<!-- grade_id: numerical_integral, points: 0, grade: false, solution: true -->
```python

def integrate_numeric(xmin, xmax, N):
    ''' 
    Numerical integral of f from xmin to xmax using Simpson's rule with 
        N panels.
    '''
    # YOUR CODE HERE
    h = (xmax-xmin)/N
    k = 0.0
    x = xmin + h 
    for i in range (1,int(N/2)+1):
        k = 4*f(x)
        x = 2*h
    x = xmin + 2*h    
    for i in range(1,(int(N/2))):
        k = 2*f(x)
        x = 2*h
    return ((f(xmin)+f(xmax)+k)*(h/3))

# It is calculating a really high number for the integral, but I was not able to spot the mistake.


```

Make sure you have implemented Simpson's rule correctly:

<!-- grade_id: cell-038d65c4eff18248, points: 3, grade: true, solution: false -->
```python
assert numpy.isclose(integrate_numeric(xmin=0, xmax=4, N=1), 1.6266126)
```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-14-637ae7f7f6d6> in <module>()
----> 1 assert numpy.isclose(integrate_numeric(xmin=0, xmax=4, N=1), 1.6266126)

AssertionError:
```

<!-- grade_id: cell-3304084098accd44, points: 3, grade: true, solution: false -->
```python
assert numpy.isclose(integrate_numeric(xmin=0, xmax=4, N=50), 1.096591)
```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-148-6323e3e60f90> in <module>()
----> 1 assert numpy.isclose(integrate_numeric(xmin=0, xmax=4, N=50), 1.096591)

AssertionError:
```

## Plotting task

** Task 1 **

There will always be some discrepancy between a numerically calculated result and an analytically derived result. Produce a log-log plot showing the fractional error between these two results as the number of panels is varied. The plot should have labels and a title.


```python
import matplotlib.pyplot as plt
x0, x1 = 0, 2  # Bounds to integrate f(x) over
panel_counts = [4, 8, 16, 32, 64, 128, 256, 512, 1024]  # Panel numbers to use
result_analytic = integrate_analytic(x0, x1)  # Define reference value from analytical solution
result_numeric = integrate_numeric(x0,x1,panel_counts)
plt.figure(figsize=(5,5))
for N in panel_counts:
    diff = (result_analytic-result_numeric)
    

plt.loglog(panel_counts, abs(result_analytic-result_numeric),'-y')
plt.xlabel('N')
plt.ylabel('Integration error')
plt.title('graph to show difference between numerical and analytical method of integration')
plt.show()
```

Output:
```
<ipython-input-12-1d73592f0014> in <module>()
      3 panel_counts = [4, 8, 16, 32, 64, 128, 256, 512, 1024]  # Panel numbers to use
      4 result_analytic = integrate_analytic(x0, x1)  # Define reference value from analytical solution
----> 5 result_numeric = integrate_numeric(x0,x1,panel_counts)
      6 plt.figure(figsize=(5,5))
      7 for N in panel_counts:

<ipython-input-9-a765e291652f> in integrate_numeric(xmin, xmax, N)
      6     '''
      7     # YOUR CODE HERE
----> 8     h = (xmax-xmin)/N
      9     k = 0.0
     10     x = xmin + h

TypeError: unsupported operand type(s) for /: 'int' and 'list'
```

What effect(s) does changing the number of panels used have
on the accuracy of the numerical method? What happens if the number of panels is taken too large?

<!-- grade_id: answer1, points: 3, grade: true, solution: true -->
YOUR ANSWER HERE

Incresing the number of panels accounts for a better aproximation of the function and therefore a reducing the integration error. If the number of panels are taken too large there can be problems with the numerical round-off, ending up with a round-off error which happenes when there is difference between the result obtained using exact arithmetic and the same result calculated when using finite precision.


If the trapezium rule was being used, how would the panel
count affect accuracy? 

<!-- grade_id: cell-eeb89d1f4669ff1d, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE 

In general the trepezium rule gives off a bigger integration error with respect to the same number of N panels compared to the Simpson rule as it approximated functions with linear segments, however it proves itself to more optimal in case the function being integrated if the function inegrated is not high order

```python

```