# <center>L2 Computational Physics</center>

## <center>Week 3: Differential Equations I</center>

```python
# usual packages to import
import numpy 
import matplotlib.pyplot as plt
%matplotlib inline
```

In this notebook, you will generate and plot the decay curve for Iodine-133 analytically and numerically. $^{133}\textrm{I}$ has a half life $t_{1/2}$ of 20.8 hours. This means that half of the nuclei will have decayed after time $t_{1/2}$. Derive the mean lifetime $\tau$ from that information.

<!-- grade_id: constants, points: 0, grade: false, solution: true -->
```python
# define a function to calculate the mean lifetime from the half life
def meanLifetime(halfLife):
    return halfLife / numpy.log(2)

T_HALF = 20.8
TAU = meanLifetime(T_HALF)

```

Check your average lifetime:

<!-- grade_id: constants_check, points: 1, grade: true, solution: false -->
```python
# this test is worth 1 mark
assert numpy.isclose(TAU, 30.0080568505)         
```

### The Decay Equation

Implement the function `f_rad` such that the differential equation 

$$ \frac{dN}{dt} = f_{rad}(N,t)$$

describes the radioactive decay process.

- *Your function should return values using hours as the time unit.*
- *The function should use the constant* `TAU`.

<!-- grade_id: decay_function, points: 0, grade: false, solution: true -->
```python
def f_rad(N, t):
    return (- N) / TAU
```

Make sure your function works:

<!-- grade_id: decay_test, points: 1, grade: true, solution: false -->
```python
# this test cell is worth 1 mark
assert numpy.isclose(f_rad(1000, 0), -33.324383681)           
```

Solve this first order, ordinary differential equation analytically. Implement this function below, naming it `analytic`. The function should take an initial number of atoms `N0` at time `t=0`, and a time argument. The function should return nuclei count at the time argument. Make sure the function also works for numpy arrays.

<!-- grade_id: analytic_function, points: 0, grade: false, solution: true -->
```python
def analytic(N0, t):
    y = N0 * numpy.exp((- t) / TAU)
    return y
```

Check your answer for a single time:

<!-- grade_id: analytic_test, points: 1, grade: true, solution: false -->
```python
# this test is worth 1 mark
assert numpy.isclose(analytic(1000, 41.6), 250.0)           
```

Output:
```
250.0
```

<!-- grade_id: analytic_test_array, points: 2, grade: true, solution: false -->
```python
# this test is worth 1 mark
assert numpy.isclose(analytic(1000, numpy.arange(0, 60, 6)), 
                     [1000.        ,  818.77471839,  670.39203948,  548.90005334,
                       449.4254866 ,  367.97822623,  301.29126855,  246.68967356,
                       201.983268  ,  165.37879338]).all()

```

Output:
```
[1000.          818.77471839  670.39203948  548.90005334  449.4254866
  367.97822623  301.29126855  246.68967356  201.983268    165.37879338]
```

## Numerically Solving the ODE

We now wish to solve our differential equation numerically. We shall do this using Euler's and RK4 methods.

### Euler's Method

Create a function which takes as its arguments the initial number of atoms, `n0`, the initial time `t0`, the time step, `dt`, and the number of steps to perform, `n_steps`.  This function should return an array of the number of counts at each time step using Euler's method. This array should contain the initial and final values, so the array length should be `n_steps+1` 

<!-- grade_id: euler_function, points: 0, grade: false, solution: true -->
```python
def solve_euler(f, n0, t0, dt, n_panels):
    eulerarray = []
    n = n0
    p = True
    for i in range(n_panels+1):
        if n == n0 and p == True:
            eulerarray.append(n)
            p = False
        else:
            for a in range(dt):
                n += f(n, t0)
                print(n)

            eulerarray.append(n)
            
    print(eulerarray)
    return eulerarray
```

Try your solution:

<!-- grade_id: euler_test_length, points: 1, grade: true, solution: false -->
```python
# this test is worth 1 mark
assert len(solve_euler(f_rad, 1000, 0, 1, 17)) == 18
```

Output:
```
966.6756163192334
934.4617471861698
903.3213853879383
873.2187569542291
844.1192800602443
815.9895252991842
788.7971772786275
762.5109974966888
737.1007884553051
712.5373589694249
688.7924906322477
665.8389053979878
643.6502342449236
622.2009868827305
601.4665224692988
581.4230213033954
562.0474574606505
[1000, 966.6756163192334, 934.4617471861698, 903.3213853879383, 873.2187569542291, 844.1192800602443, 815.9895252991842, 788.7971772786275, 762.5109974966888, 737.1007884553051, 712.5373589694249, 688.7924906322477, 665.8389053979878, 643.6502342449236, 622.2009868827305, 601.4665224692988, 581.4230213033954, 562.0474574606505]
```

<!-- grade_id: euler_test_one, points: 2, grade: true, solution: false -->
```python
# this test is worth 2 marks
assert numpy.isclose(solve_euler(f_rad, 1000, 0, 6, 1), [1000.,  800.05369792]).all()
```

Output:
```
966.6756163192334
934.4617471861698
903.3213853879383
873.2187569542291
844.1192800602443
815.9895252991842
[1000, 815.9895252991842]
AssertionError                            Traceback (most recent call last)
<ipython-input-90-b79d7ba7b283> in <module>()
      1 # this test is worth 2 marks
----> 2 assert numpy.isclose(solve_euler(f_rad, 1000, 0, 6, 1), [1000.,  800.05369792]).all()

AssertionError:
```

<!-- grade_id: euler_test_many, points: 2, grade: true, solution: false -->
```python
# this test is worth 2 mark
assert numpy.isclose(solve_euler(f_rad, 1000, 0, 6, 10), [1000.        ,  800.05369792,  640.08591955,  512.10310692,
                                                409.7099844 ,  327.7899881 ,  262.24959212,  209.81375595,
                                                167.86227132,  134.29883091,  107.4462763 ]).all()
```

Output:
```
966.6756163192334
934.4617471861698
903.3213853879383
873.2187569542291
844.1192800602443
815.9895252991842
788.7971772786275
762.5109974966888
737.1007884553051
712.5373589694249
688.7924906322477
665.8389053979878
643.6502342449236
622.2009868827305
601.4665224692988
581.4230213033954
562.0474574606505
543.3175723414325
525.211849100224
507.7094879271232
490.7903821530742
474.4350951513751
458.62483800892966
443.3414479415907
428.5673674287985
414.2856240435451
400.47981095449217
387.1340680778438
374.23306385732167
361.76197765131144
349.70648270694625
338.0527297015686
326.78733083266303
315.89734443798176
305.3702601281952
295.1939844149877
285.35682681808845
275.8474864352764
266.6550389599322
257.76892413122164
249.1789336024944
240.8751992139606
232.8481816561735
225.08865951129434
217.58771865955052
210.33674203871695
203.3273997448563
196.5516394629461
190.00167721639917
183.66998842485071
177.54929925993903
171.63257828914956
165.91302839812272
160.38407898214575
155.0393783
```

### RK 4 method

Implement the RK4 method in the `solve_RK4` function. The arguments are the same as for `solve_euler`.

<!-- grade_id: RK4, points: 0, grade: false, solution: true -->
```python
def solve_RK4(f, n0, t0, dt, nsteps):
    # YOUR CODE HERE

```

<!-- grade_id: RK4-test-1, points: 1, grade: true, solution: false -->
```python
# This checks that we return an array of the right length
# this test is worth 1 mark
assert len(solve_RK4(f_rad, 1000, 0, 1, 17)) == 18
```

<!-- grade_id: RK4-test-2, points: 2, grade: true, solution: false -->
```python
# This checks that a single step is working
# this test is worth 2 mark
assert numpy.isclose(solve_RK4(f_rad, 1000,0, 6, 1), [1000.,  818.7773]).all()
```

<!-- grade_id: RK4-test-3, points: 2, grade: true, solution: false -->
```python
# This checks multiple steps
# this test is worth 2 marks
assert numpy.isclose(solve_RK4(f_rad, 1000, 0, 6, 10), [
    1000.,
    818.77729521,  
    670.39625915,  
    548.90523578,
    449.43114428,  
    367.9840167,  
    301.29695787,  
    246.69510822, 
    201.98835345,  
    165.3834777,  
    135.41223655]).all()
```

## Plotting task

**Task 1: **

Create a plot to show that the RK4 method has an error that scales better with the number of steps than the Euler method. (click on the "+" button to create new cells.)       [task worth 5 marks]


```python

```