# Multi-dimensional Monte Carlo integration

In this notebook you will implement a Monte Carlo integration function to estimate the size of the unit ball in different dimensions.

```python
import numpy
import matplotlib.pyplot as plt
```

Implement the function `integrate` that gives the Monte Carlo estimate for the volume of a unit ball in `dim` dimensions. To do so we choose random numbers in the `dim`-dimensional cube of sides $[-1,+1]$ that contains the unit ball. Use the random numbers provided, they form a `npoints` dimensional array of `dim` coordinates in the cube.

<!-- grade_id: integrate, points: 0, grade: false, solution: true -->
```python
def integrate(npoints, dim):
    # the random numbers
    rs = numpy.random.uniform(-1, 1, size=(npoints,dim)) 
    # YOUR CODE HERE
```

Check your answer:

<!-- grade_id: d2test, points: 1, grade: true, solution: false -->
```python
# this test is worth 1 mark
numpy.random.seed(1234)
assert numpy.isclose(integrate(1000000, 2), 3.140204)
```

<!-- grade_id: d3test, points: 1, grade: true, solution: false -->
```python
# this test is worth 1 mark
numpy.random.seed(1234)
assert numpy.isclose(integrate(100000, 3), 4.20312)
```

<!-- grade_id: d4test, points: 1, grade: true, solution: false -->
```python
# this test is worth 1 mark
numpy.random.seed(1234)
assert numpy.isclose(integrate(100000, 4), 4.97312)
```

<!-- grade_id: d5test, points: 1, grade: true, solution: false -->
```python
# this test is worth 1 mark
numpy.random.seed(1234)
assert numpy.isclose(integrate(100000, 5), 5.31616)
```

<!-- grade_id: d6test, points: 1, grade: true, solution: false -->
```python
# this test is worth 1 mark
numpy.random.seed(1234)
assert numpy.isclose(integrate(100000, 6), 5.20448)
```

Use your function to show that the integration error scaling is independent of the dimension and that the error scales as $1/\sqrt{N}$. Use the provided number of points and dimensions. Your plot should have labels, a title and a legend.

| dimension | volume of the unit ball |
| --------- | ----------------------- |
| 2         |  $\pi$                  |
|    3      | $4/3\pi$                |
|    4      | $1/2\pi^2$              |
|    5      | $8/15\pi^2$             |
|    6      | $1/6\pi^3$              |


Your plot should be generated in in less than 30 seconds, otherwise the plots will not be generated for the markers. [5 marks]


```python
ns = [2**ii for ii in range(4,20)]
dimensions = range(2,7)

```