# Assignment 1

First we need to import a few modules, execute the cell below:

```python
import numpy 
import matplotlib.pyplot as plt
%matplotlib inline
```

Define a function `f` that is equivalent to $f(x)=\cos(x)$. It should work both for single arguments and `numpy` arrays.

<!-- grade_id: f_definition, points: 0, grade: false, solution: true -->
```python

def f(x):
    f = numpy.cos(x)
    print(f)
    return f;
assert numpy.isclose(f(0.5),0.87758256189)
```

Output:
```
0.8775825618903728
```

You can use the following cell to check that your code works for scalar arguments:

<!-- grade_id: f_definition_test, points: 1, grade: true, solution: false -->
```python
assert numpy.isclose(f(0.5),0.87758256189)
```

Output:
```
0.8775825618903728
```

Make sure your code also works with a numpy array:

<!-- grade_id: f_definition_nparray_test, points: 2, grade: true, solution: false -->
```python
xs = numpy.arange(start=0.1, stop=0.7, step=0.1)
assert numpy.isclose(
            f(xs), 
            [0.99500417, 0.98006658, 0.95533649, 0.92106099, 0.87758256, 0.82533561]
        ).all()
```

Output:
```
[0.99500417 0.98006658 0.95533649 0.92106099 0.87758256 0.82533561]
```

Define the analytic derivative of the function`f` here (again, it should work both for single numbers and numpy arrays).

<!-- grade_id: analytic_derivative, points: 0, grade: false, solution: true -->
```python
import numpy as np

def df_analytic(x):   
    f =-numpy.sin(x)
    print (df_analytic)
    return f;

```

Output:
```
<function df_analytic at 0x7f464ee11730>
```

<!-- grade_id: g_definition_test, points: 1, grade: true, solution: false -->
```python
assert numpy.isclose(df_analytic(0.5), -0.4794255386)
```

Output:
```
<function df_analytic at 0x7f464ee11730>
```

<!-- grade_id: g_definition_nparray_test, points: 2, grade: true, solution: false -->
```python
xs = numpy.arange(start=0.1, stop=0.7, step=0.1)
assert numpy.isclose(
            df_analytic(xs),
            [-0.09983342, -0.19866933, -0.29552021, -0.38941834, -0.47942554,-0.56464247]
        ).all()
```

Output:
```
<function df_analytic at 0x7f464ee11730>
```

Now define the numerical derivative using the backward difference method. The function `backward_difference` takes three arguments, `f`, the function to calculate the derivative for, `x` the position to calculate the derivative at and `dx` the interval length.

<!-- grade_id: backward_difference, points: 0, grade: false, solution: true -->
```python
import numpy as np

def f(x):
    f = np.cos(x)
    print(f)
    return;
def backward_difference(f, x, dx):
    backward_difference = (f(x)-f(x-dx))/dx
    print(backward_difference)
    return;
backward_difference(f, 0.1, 0.0001) 
```

Output:
```
0.9950041652780258
0.995014143644653
<ipython-input-20-e01be0fef5a9> in <module>()
      9     print(backward_difference)
     10     return;
---> 11 backward_difference(f, 0.1, 0.0001)

<ipython-input-20-e01be0fef5a9> in backward_difference(f, x, dx)
      6     return;
      7 def backward_difference(f, x, dx):
----> 8     backward_difference = (f(x)-f(x-dx))/dx
      9     print(backward_difference)
     10     return;

TypeError: unsupported operand type(s) for -: 'NoneType' and 'NoneType'
```

<!-- grade_id: backward_difference_test, points: 2, grade: true, solution: false -->
```python
assert numpy.isclose(
    backward_difference(f, 0.1, 0.0001),
    -0.09978366627194646
)
```

**Task 1:**

Here is a skeleton code to plot the difference between the analytical derivative and the numerical implementation. Modify and expand it to provide a plot with three curves for the difference for a case where $dx$ is too large, about right and too small. The plot should have a legend, labelled axes and a title.

You can get help about plotting on the Python handbook that you can find at:

https://dmaitre.phyip3.dur.ac.uk/notes/l1python/part-5/


```python
import numpy 
import matplotlib.pyplot as plt

def f(x):
    f = numpy.cos(x)
    print(f)
    return;

fig = plt.figure()
xs = numpy.linspace(-2*numpy.pi,2*numpy.pi,100)

df_dx_1 = backward_difference(f, xs, dx=1e-4)

df_dx_2 = backward_difference(f, xs, dx=1e-6)

plt.plot(xs, df_dx_1 - df_dx_analytical)

plt.plot(xs, df_dx_2 - df_dx_analytical) 

df_dx_analytical = df_analytic(xs)

plt.xlabel("x")
plt.ylabel("derivative")
fig.suptitle("the difference between the analytical derivative and the numerical implementation")


plt.figure(figsize=(8, 4))

plt.show()

```

Output:
```
[ 1.          0.99195481  0.9679487   0.92836793  0.87384938  0.80527026
  0.72373404  0.63055267  0.52722547  0.41541501  0.29692038  0.17364818
  0.04758192 -0.07924996 -0.20480667 -0.32706796 -0.44406661 -0.55392006
 -0.65486073 -0.74526445 -0.82367658 -0.88883545 -0.93969262 -0.97542979
 -0.99547192 -0.99949654 -0.98743889 -0.95949297 -0.91610846 -0.85798341
 -0.78605309 -0.70147489 -0.60560969 -0.5        -0.38634513 -0.26647381
 -0.14231484 -0.01586596  0.1108382   0.23575894  0.35688622  0.47227107
  0.58005691  0.67850941  0.76604444  0.84125353  0.90292654  0.95007112
  0.9819287   0.99798668  0.99798668  0.9819287   0.95007112  0.90292654
  0.84125353  0.76604444  0.67850941  0.58005691  0.47227107  0.35688622
  0.23575894  0.1108382  -0.01586596 -0.14231484 -0.26647381 -0.38634513
 -0.5        -0.60560969 -0.70147489 -0.78605309 -0.85798341 -0.91610846
 -0.95949297 -0.98743889 -0.99949654 -0.99547192 -0.97542979 -0.93969262
 -0.88883545 -0.82367658 -0.74526445 -0.65486073 -0
```

**Task 2: ** 

Describe in the text box below why the numerical method looses accuracy in the case where
$\Delta x$ is too small and the case where $\Delta x$ is too large. (Double-click on "YOUR ANSWER HERE" to edit.)


<!-- grade_id: textAnswer, points: 4, grade: true, solution: true -->
The numerical method looses accuracy in the case where $\Delta x$ is small and in the case where it is large because the computer can only store a finite amount of decimal numbers. For this reason, when it performs calculations, numbers past a certain decimal place are discarded by the computer so part of the information is lost, meaning that the computer sometimes 'makes up' decimal places if that information is needed later on in the calculation.
To illustrate, imagine a computer has 5 digits precision for x, so $x = 1.23456⋅10^{-3}$
Doing $1+x$ gives $1.00123$ (represented as $1.00123⋅10^0$) and now substract $1$, we get $1.23⋆⋆⋆ ⋅ 10^{-3}$, with the computer having to make up the lost information

```python

```