# <center>Computational Physics</center>
---

## Week 2: Numerical Integration

```python
import numpy 
import matplotlib.pyplot as plt
%matplotlib inline
```

Define the function `f`, such that $\textrm{f}(x) \equiv x^{2}\sin(x)$. This is the function that we will be integrating.

<!-- grade_id: integrand, points: 0, grade: false, solution: true -->
```python
def f(x):
    '''Function equivalent to x^2 sin(x).'''
    return (x**2) * numpy.sin(x)
```

Ensure your function works with numpy arrays:

<!-- grade_id: cell-653afd49e7782cc4, points: 1, grade: true, solution: false -->
```python
xs=numpy.arange(0, 1, step=0.1)
assert numpy.isclose(f(xs), 
                     [0., 0.00099833, 0.00794677, 0.02659682, 0.06230693,
                      0.11985638, 0.20327129, 0.31566667, 0.4591079 , 0.6344948 ]).all()
```

Derive the indefinite integral of $f(x)$ nalytically. Call this function $g(x)$ and implement it below. Set the constant of integration such that $g(0)=0$.

<!-- grade_id: integral, points: 0, grade: false, solution: true -->
```python
def g(x):
    '''Analytical integral of f(x).'''
    return -x**2 * numpy.cos(x) + 2 * (x * numpy.sin(x) + numpy.cos(x)) - 2
```

Check your solution with the same numpy array:

<!-- grade_id: cell-648a123ee323b45e, points: 1, grade: true, solution: false -->
```python
assert g(0) == 0.
```

<!-- grade_id: cell-f8eaec03e0540a78, points: 1, grade: true, solution: false -->
```python
assert numpy.isclose(g(xs), 
                     [0., 0.00002497, 0.00039822, 0.00200482, 0.0062869, 
                      0.01519502, 0.03112138, 0.05681646, 0.09529087, 0.1497043 ]).all()
```

Now, using the analytically derived indefinite integral, $g(x)$, define a function which calculates the definite integral of $f(x)$ over the interval $(x_{min},~x_{max})$.

<!-- grade_id: cell-6f8d8cc5f5c7564c, points: 0, grade: false, solution: true -->
```python
def integrate_analytic(xmin, xmax):
    '''Analytical integral of f(x) from xmin to xmax.'''
    return g(xmax) - g(xmin)
```

Check your analytic function:

<!-- grade_id: analytical_integral, points: 1, grade: true, solution: false -->
```python
assert numpy.isclose(integrate_analytic(xmin=0, xmax=4), 1.096591)
```

## Numerical implementation

Create a function which calculates the definite integral of the function $f(x)$ over the interval $(x_{min},~x_{max})$ using Simpson's rule with $N$ panels.

<!-- grade_id: numerical_integral, points: 0, grade: false, solution: true -->
```python
def integrate_numeric(xmin, xmax, N):
    ''' 
    Numerical integral of f from xmin to xmax using Simpson's rule with 
        N panels.
    '''
    xsum = 0
    msum = 0
    h = (xmax-xmin)/N
    
    for i in range(0, N):
        xsum += f(xmin + i*h)
        print (xsum)
        
    for i in range(0,N-1):
        msum += f(xmin + (h/2) + i*h)    
        print (msum)
    
    I = (h/6) * (f(xmin) + 4*(msum) + 2*(xsum) + f(xmax))
    return I
```

Make sure you have implemented Simpson's rule correctly:

<!-- grade_id: cell-038d65c4eff18248, points: 3, grade: true, solution: false -->
```python
assert numpy.isclose(integrate_numeric(xmin=0, xmax=4, N=1), 1.6266126)
```

Output:
```
0.0
AssertionError                            Traceback (most recent call last)
<ipython-input-48-637ae7f7f6d6> in <module>()
----> 1 assert numpy.isclose(integrate_numeric(xmin=0, xmax=4, N=1), 1.6266126)

AssertionError:
```

<!-- grade_id: cell-3304084098accd44, points: 3, grade: true, solution: false -->
```python
assert numpy.isclose(integrate_numeric(xmin=0, xmax=4, N=50), 1.096591)
```

Output:
```
0.0
0.0005114540414027053
0.004590000130727403
0.018281671412930356
0.05049328722002082
0.11280022198940491
0.21919414403416257
0.38577413570215163
0.6303853884841875
0.9722104024341991
1.4313183006098937
2.0281784884292557
2.7831454377754556
3.7159218499618505
4.845007844628059
6.187144128420865
7.756757313918731
9.565415682583092
11.621303713393374
13.928723630034561
16.487632053820814
19.293219587704662
22.33554080047015
25.599201632243876
29.063110706437456
32.70030041374018
36.4778229362815
40.35672561089052
44.292109196594105
48.23327172094253
52.123939640917165
55.902587076125414
59.50284286429423
62.85398416171911
65.88151427462057
68.50782137184287
70.65291370567198
72.2352259664541
73.1724904287976
73.3826656229057
72.78491439420729
71.30062240775652
68.85444742021517
65.37538899041431
60.797867737664816
55.06280279256353
48.11867572462938
39.9225689804195
30.44116672958938
19.651705998436334
6.398293469861466e-05
0.0017878387196590534
0.009734611951461502
0.03140089479888801
```

## Plotting task

** Task 1 **

There will always be some discrepancy between a numerically calculated result and an analytically derived result. Produce a log-log plot showing the fractional error between these two results as the number of panels is varied. The plot should have labels and a title.


```python
x0, x1 = 0, 2  # Bounds to integrate f(x) over
panel_counts = [4, 8, 16, 32, 64, 128, 256, 512, 1024]  # Panel numbers to use
result_analytic = integrate_analytic(x0, x1)  # Define reference value from analytical solution
```

What effect(s) does changing the number of panels used have
on the accuracy of the numerical method? What happens if the number of panels is taken too large?

<!-- grade_id: answer1, points: 3, grade: true, solution: true -->
YOUR ANSWER HERE

If the trapezium rule was being used, how would the panel
count affect accuracy? 

<!-- grade_id: cell-eeb89d1f4669ff1d, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE