# <center>Computational Physics</center>
---

## Week 2: Numerical Integration

```python
import numpy 
import matplotlib.pyplot as plt
%matplotlib inline
```

Define the function `f`, such that $\textrm{f}(x) \equiv x^{2}\sin(x)$. This is the function that we will be integrating.

<!-- grade_id: integrand, points: 0, grade: false, solution: true -->
```python
def f(x):
    '''Function equivalent to x^2 sin(x).'''
    y1=(x**2)*numpy.sin(x)
    print(y1)
    return y1;
```

Ensure your function works with numpy arrays:

<!-- grade_id: cell-653afd49e7782cc4, points: 1, grade: true, solution: false -->
```python
xs=numpy.arange(0, 1, step=0.1)
assert numpy.isclose(f(xs), 
                     [0., 0.00099833, 0.00794677, 0.02659682, 0.06230693,
                      0.11985638, 0.20327129, 0.31566667, 0.4591079 , 0.6344948 ]).all()
```

Output:
```
[0.         0.00099833 0.00794677 0.02659682 0.06230693 0.11985638
 0.20327129 0.31566667 0.4591079  0.6344948 ]
```

Derive the indefinite integral of $f(x)$ nalytically. Call this function $g(x)$ and implement it below. Set the constant of integration such that $g(0)=0$.

<!-- grade_id: integral, points: 0, grade: false, solution: true -->
```python
def g(x):
    '''Analytical integral of f(x).'''
    y2=-(x**2)*numpy.cos(x)+2*x*numpy.sin(x)+2*numpy.cos(x)-2
    print(y2)
    return y2;
```

Check your solution with the same numpy array:

<!-- grade_id: cell-648a123ee323b45e, points: 1, grade: true, solution: false -->
```python
assert g(0) == 0.
```

Output:
```
0.0
```

<!-- grade_id: cell-f8eaec03e0540a78, points: 1, grade: true, solution: false -->
```python
assert numpy.isclose(g(xs), 
                     [0., 0.00002497, 0.00039822, 0.00200482, 0.0062869, 
                      0.01519502, 0.03112138, 0.05681646, 0.09529087, 0.1497043 ]).all()
```

Output:
```
[0.00000000e+00 2.49722326e-05 3.98224887e-04 2.00481823e-03
 6.28690281e-03 1.51950219e-02 3.11213765e-02 5.68164649e-02
 9.52908702e-02 1.49704300e-01]
```

Now, using the analytically derived indefinite integral, $g(x)$, define a function which calculates the definite integral of $f(x)$ over the interval $(x_{min},~x_{max})$.

<!-- grade_id: cell-6f8d8cc5f5c7564c, points: 0, grade: false, solution: true -->
```python
def integrate_analytic(xmin, xmax):
    '''Analytical integral of f(x) from xmin to xmax.'''
    y3=g(xmax)-g(xmin)
    print(y3)
    return y3;


```

Check your analytic function:

<!-- grade_id: analytical_integral, points: 1, grade: true, solution: false -->
```python
assert numpy.isclose(integrate_analytic(xmin=0, xmax=4), 1.096591)
```

Output:
```
1.0965907296271418
0.0
1.0965907296271418
```

## Numerical implementation

Create a function which calculates the definite integral of the function $f(x)$ over the interval $(x_{min},~x_{max})$ using Simpson's rule with $N$ panels.

<!-- grade_id: numerical_integral, points: 0, grade: false, solution: true -->
```python
def integrate_numeric(xmin, xmax, N):
    ''' 
    Numerical integral of f from xmin to xmax using Simpson's rule with 
        N panels.
    '''
    h=(xmax-xmin)/N
    y4 = f(x)
    k=0.0
    x=a
    for i in range(1,n/2):
        x += 2*h
        k += 4*f(x)
    for i in range(2,(n/2)-1):
        x += 2*h
        k += 2*f(x)
    return (h/3)*(f(a)+f(b)+k)

def function(x): return x

print simpson(function, 0.0, 1.0, 100)
```

Make sure you have implemented Simpson's rule correctly:

<!-- grade_id: cell-038d65c4eff18248, points: 3, grade: true, solution: false -->
```python
assert numpy.isclose(integrate_numeric(xmin=0, xmax=4, N=1), 1.6266126)
```

Output:
```
[  0.         -12.10883992]
0.0
AssertionError                            Traceback (most recent call last)
<ipython-input-45-637ae7f7f6d6> in <module>()
----> 1 assert numpy.isclose(integrate_numeric(xmin=0, xmax=4, N=1), 1.6266126)

AssertionError:
```

<!-- grade_id: cell-3304084098accd44, points: 3, grade: true, solution: false -->
```python
assert numpy.isclose(integrate_numeric(xmin=0, xmax=4, N=50), 1.096591)
```

Output:
```
[ 0.00000000e+00  5.11454041e-04  4.07854609e-03  1.36916713e-02
  3.22116158e-02  6.23069348e-02  1.06393922e-01  1.66579992e-01
  2.44611253e-01  3.41825014e-01  4.59107898e-01  5.96860188e-01
  7.54966949e-01  9.32776412e-01  1.12908599e+00  1.34213628e+00
  1.56961319e+00  1.80865837e+00  2.05588803e+00  2.30741992e+00
  2.55890842e+00  2.80558753e+00  3.04232121e+00  3.26366083e+00
  3.46390907e+00  3.63718971e+00  3.77752252e+00  3.87890267e+00
  3.93538359e+00  3.94116252e+00  3.89066792e+00  3.77864744e+00
  3.60025579e+00  3.35114130e+00  3.02753011e+00  2.62630710e+00
  2.14509233e+00  1.58231226e+00  9.37264462e-01  2.10175194e-01
 -5.97751229e-01 -1.48429199e+00 -2.44617499e+00 -3.47905843e+00
 -4.57752125e+00 -5.73506495e+00 -6.94412707e+00 -8.19610674e+00
 -9.48140225e+00 -1.07894607e+01 -1.21088399e+01]
1.096594986172562
```

## Plotting task

** Task 1 **

There will always be some discrepancy between a numerically calculated result and an analytically derived result. Produce a log-log plot showing the fractional error between these two results as the number of panels is varied. The plot should have labels and a title.


```python
x0, x1 = 0, 2  # Bounds to integrate f(x) over
panel_counts = [4, 8, 16, 32, 64, 128, 256, 512, 1024]  # Panel numbers to use
result_analytic = integrate_analytic(x0, x1)  # Define reference value from analytical solution
```

What effect(s) does changing the number of panels used have
on the accuracy of the numerical method? What happens if the number of panels is taken too large?

<!-- grade_id: answer1, points: 3, grade: true, solution: true -->
YOUR ANSWER HERE

If the trapezium rule was being used, how would the panel
count affect accuracy? 

<!-- grade_id: cell-eeb89d1f4669ff1d, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE