# <center>Computational Physics</center>
---

## Week 2: Numerical Integration

```python
import numpy 
import matplotlib.pyplot as plt
%matplotlib inline
```

Define the function `f`, such that $\textrm{f}(x) \equiv x^{2}\sin(x)$. This is the function that we will be integrating.

<!-- grade_id: integrand, points: 0, grade: false, solution: true -->
```python
def f(x):
    value = (x**2)*(numpy.sin(x))
    print(value)
    return value
```

Ensure your function works with numpy arrays:

<!-- grade_id: cell-653afd49e7782cc4, points: 1, grade: true, solution: false -->
```python
xs=numpy.arange(0, 1, step=0.1)
assert numpy.isclose(f(xs), 
                     [0., 0.00099833, 0.00794677, 0.02659682, 0.06230693,
                      0.11985638, 0.20327129, 0.31566667, 0.4591079 , 0.6344948 ]).all()
```

Output:
```
[0.         0.00099833 0.00794677 0.02659682 0.06230693 0.11985638
 0.20327129 0.31566667 0.4591079  0.6344948 ]
```

Derive the indefinite integral of $f(x)$ nalytically. Call this function $g(x)$ and implement it below. Set the constant of integration such that $g(0)=0$.

<!-- grade_id: integral, points: 0, grade: false, solution: true -->
```python
def g(x):
    value2 = (2*x)*(numpy.sin(x))-(x**2-2)*(numpy.cos(x)) - 2
    print(value2)
    return value2

```

Check your solution with the same numpy array:

<!-- grade_id: cell-648a123ee323b45e, points: 1, grade: true, solution: false -->
```python
assert g(0) == 0.
```

Output:
```
0.0
```

<!-- grade_id: cell-f8eaec03e0540a78, points: 1, grade: true, solution: false -->
```python
assert numpy.isclose(g(xs), 
                     [0., 0.00002497, 0.00039822, 0.00200482, 0.0062869, 
                      0.01519502, 0.03112138, 0.05681646, 0.09529087, 0.1497043 ]).all()
```

Output:
```
[0.00000000e+00 2.49722326e-05 3.98224887e-04 2.00481823e-03
 6.28690281e-03 1.51950219e-02 3.11213765e-02 5.68164649e-02
 9.52908702e-02 1.49704300e-01]
```

Now, using the analytically derived indefinite integral, $g(x)$, define a function which calculates the definite integral of $f(x)$ over the interval $(x_{min},~x_{max})$.

<!-- grade_id: cell-6f8d8cc5f5c7564c, points: 0, grade: false, solution: true -->
```python
def integrate_analytic(xmin, xmax):
    value3 = g(xmax) - g(xmin)
    print(value3)
    return value3
```

Check your analytic function:

<!-- grade_id: analytical_integral, points: 1, grade: true, solution: false -->
```python
assert numpy.isclose(integrate_analytic(xmin=0, xmax=4), 1.096591)
```

Output:
```
1.0965907296271418
0.0
1.0965907296271418
```

## Numerical implementation

Create a function which calculates the definite integral of the function $f(x)$ over the interval $(x_{min},~x_{max})$ using Simpson's rule with $N$ panels.

<!-- grade_id: numerical_integral, points: 0, grade: false, solution: true -->
```python
def integrate_numeric(xmin, xmax, N):
    h = (xmax-xmin)/N
    k = h/2
    
    total_first = 0
    total_second = 0
    
    for i in range (0,N):
        total_first = total_first + 4*f(xmin + (2*i + 1)*k)
        print(total_first)
    
    for i in range (0,N-1):
        total_second = total_second + 2*f(xmin + 2*(i + 1)*k)
        print(total_second)
   
    answer = (h/6)*(f(xmin)+ f(xmax)+ total_first + total_second)
    print(answer)
    return answer



```

Make sure you have implemented Simpson's rule correctly:

<!-- grade_id: cell-038d65c4eff18248, points: 3, grade: true, solution: false -->
```python
assert numpy.isclose(integrate_numeric(xmin=0, xmax=4, N=1), 1.6266126)
```

Output:
```
3.637189707302727
14.548758829210907
0.0
-12.108839924926851
1.6266126028560373
```

<!-- grade_id: cell-3304084098accd44, points: 3, grade: true, solution: false -->
```python
assert numpy.isclose(integrate_numeric(xmin=0, xmax=4, N=50), 1.096591)
```

Output:
```
6.398293469861466e-05
0.00025593173879445865
0.0017238557849604387
0.007151354878636214
0.00794677323180245
0.03893844780584601
0.021666282847426523
0.1256035791955521
0.04565474063245165
0.3082225417253587
0.08246188043677752
0.6380700634724688
0.13435638927294644
1.1754956205642544
0.20327129042221273
1.9885807822531052
0.29075389430613996
3.151596359477665
0.39792102669585444
4.743280466261083
0.5254201854514382
6.844961208066835
0.6733972111987803
9.538550052861957
0.8414709848078965
12.904433992093544
1.028715585950603
17.019296335895955
1.2336502631632893
21.953897388549112
1.454237477553611
27.770847298763556
1.6878891904460656
34.52240406054782
1.9314814707773824
42.24832994365735
2.1813774018752468
50.97383955115834
2.433458170351646
60.707672232564924
2.68316212319605
71.44032072534912
2.9255314837483084
83.14244666034236
3.1552663240453525
95.76351195652377
3.366785301033527
109.23065316065788
3.554292578259827
123.44782347369718
3.7118502738075843
138.29522456892752
3.83345
```

## Plotting task

** Task 1 **

There will always be some discrepancy between a numerically calculated result and an analytically derived result. Produce a log-log plot showing the fractional error between these two results as the number of panels is varied. The plot should have labels and a title.


```python
x0, x1 = 0, 2  # Bounds to integrate f(x) over
panel_counts = [4, 8, 16, 32, 64, 128, 256, 512, 1024]  # Panel numbers to use
result_analytic = integrate_analytic(x0, x1)  # Define reference value from analytical solution
```

Output:
```
2.469483380397012
0.0
2.469483380397012
```

What effect(s) does changing the number of panels used have
on the accuracy of the numerical method? What happens if the number of panels is taken too large?

<!-- grade_id: answer1, points: 3, grade: true, solution: true -->
YOUR ANSWER HERE

If the trapezium rule was being used, how would the panel
count affect accuracy? 

<!-- grade_id: cell-eeb89d1f4669ff1d, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE