# <center>L2 Computational Physics</center>

## <center>Week 3: Differential Equations I</center>

```python
# usual packages to import
import numpy 
import matplotlib.pyplot as plt
%matplotlib inline
```

In this notebook, you will generate and plot the decay curve for Iodine-133 analytically and numerically. $^{133}\textrm{I}$ has a half life $t_{1/2}$ of 20.8 hours. This means that half of the nuclei will have decayed after time $t_{1/2}$. Derive the mean lifetime $\tau$ from that information.

<!-- grade_id: constants, points: 0, grade: false, solution: true -->
```python
# define a function to calculate the mean lifetime from the half life
def meanLifetime(halfLife):
    # YOUR CODE HERE
    return halfLife/numpy.log(2)

T_HALF = 20.8
TAU = meanLifetime(T_HALF)
print(TAU)
```

Output:
```
30.00805685049044
```

Check your average lifetime:

<!-- grade_id: constants_check, points: 1, grade: true, solution: false -->
```python
# this test is worth 1 mark
assert numpy.isclose(TAU, 30.0080568505)         
```

### The Decay Equation

Implement the function `f_rad` such that the differential equation 

$$ \frac{dN}{dt} = f_{rad}(N,t)$$

describes the radioactive decay process.

- *Your function should return values using hours as the time unit.*
- *The function should use the constant* `TAU`.

<!-- grade_id: decay_function, points: 0, grade: false, solution: true -->
```python
def f_rad(N, t):
    # YOUR CODE HERE
    return -N/TAU
```

Make sure your function works:

<!-- grade_id: decay_test, points: 1, grade: true, solution: false -->
```python
# this test cell is worth 1 mark
assert numpy.isclose(f_rad(1000, 0), -33.324383681)           
```

Solve this first order, ordinary differential equation analytically. Implement this function below, naming it `analytic`. The function should take an initial number of atoms `N0` at time `t=0`, and a time argument. The function should return nuclei count at the time argument. Make sure the function also works for numpy arrays.

<!-- grade_id: analytic_function, points: 0, grade: false, solution: true -->
```python
def analytic(N0, t):
    # YOUR CODE HERE
    return N0*numpy.exp(-t/TAU)
```

Check your answer for a single time:

<!-- grade_id: analytic_test, points: 1, grade: true, solution: false -->
```python
# this test is worth 1 mark
assert numpy.isclose(analytic(1000, 41.6), 250.0)           
```

<!-- grade_id: analytic_test_array, points: 2, grade: true, solution: false -->
```python
# this test is worth 1 mark
assert numpy.isclose(analytic(1000, numpy.arange(0, 60, 6)), 
                     [1000.        ,  818.77471839,  670.39203948,  548.90005334,
                       449.4254866 ,  367.97822623,  301.29126855,  246.68967356,
                       201.983268  ,  165.37879338]).all()

```

## Numerically Solving the ODE

We now wish to solve our differential equation numerically. We shall do this using Euler's and RK4 methods.

### Euler's Method

Create a function which takes as its arguments the initial number of atoms, `n0`, the initial time `t0`, the time step, `dt`, and the number of steps to perform, `n_steps`.  This function should return an array of the number of counts at each time step using Euler's method. This array should contain the initial and final values, so the array length should be `n_steps+1` 

<!-- grade_id: euler_function, points: 0, grade: false, solution: true -->
```python
def solve_euler(f, n0, t0, dt, n_steps):
    # YOUR CODE HERE
    counts = []
    n = n0
    for i in range (0,n_steps+1):
        counts.append(n)
        slope = f(n, t0+(i*dt))
        displacement = (slope*dt)
        n = n+displacement
    array = numpy.array(counts)
    return array
```

Try your solution:

<!-- grade_id: euler_test_length, points: 1, grade: true, solution: false -->
```python
# this test is worth 1 mark
assert len(solve_euler(f_rad, 1000, 0, 1, 17)) == 18
```

<!-- grade_id: euler_test_one, points: 2, grade: true, solution: false -->
```python
# this test is worth 2 marks
assert numpy.isclose(solve_euler(f_rad, 1000, 0, 6, 1), [1000.,  800.05369792]).all()
```

<!-- grade_id: euler_test_many, points: 2, grade: true, solution: false -->
```python
# this test is worth 2 mark
assert numpy.isclose(solve_euler(f_rad, 1000, 0, 6, 10), [1000.        ,  800.05369792,  640.08591955,  512.10310692,
                                                409.7099844 ,  327.7899881 ,  262.24959212,  209.81375595,
                                                167.86227132,  134.29883091,  107.4462763 ]).all()
```

### RK 4 method

Implement the RK4 method in the `solve_RK4` function. The arguments are the same as for `solve_euler`.

<!-- grade_id: RK4, points: 0, grade: false, solution: true -->
```python
def solve_RK4(f, n0, t0, dt, n_steps):
    # YOUR CODE HERE
    counts = []
    n = n0
    t = t0
    for i in range(0,n_steps+1):
        counts.append(n)
        k1 = f(n,t)
        n1 = (dt*k1)/2
        k2 = f(n+n1,t+(dt/2))
        n2 = (dt*k2)/2
        k3 = f(n+n2,t+(dt/2))
        n3 = dt*k3
        k4 = f(n+n3,t+dt)
        k = (1/6)*(k1+(2*k2)+(2*k3)+k4)
        n = n+(dt*k)
        t = t + dt*(i+1)
    array = numpy.array(counts)
    return array
```

<!-- grade_id: RK4-test-1, points: 1, grade: true, solution: false -->
```python
# This checks that we return an array of the right length
# this test is worth 1 mark
assert len(solve_RK4(f_rad, 1000, 0, 1, 17)) == 18
```

<!-- grade_id: RK4-test-2, points: 2, grade: true, solution: false -->
```python
# This checks that a single step is working
# this test is worth 2 mark
assert numpy.isclose(solve_RK4(f_rad, 1000,0, 6, 1), [1000.,  818.7773]).all()
```

<!-- grade_id: RK4-test-3, points: 2, grade: true, solution: false -->
```python
# This checks multiple steps
# this test is worth 2 marks
assert numpy.isclose(solve_RK4(f_rad, 1000, 0, 6, 10), [
    1000.,
    818.77729521,  
    670.39625915,  
    548.90523578,
    449.43114428,  
    367.9840167,  
    301.29695787,  
    246.69510822, 
    201.98835345,  
    165.3834777,  
    135.41223655]).all()
```

## Plotting task

**Task 1: **

Create a plot to show that the RK4 method has an error that scales better with the number of steps than the Euler method. (click on the "+" button to create new cells.)       [task worth 5 marks]


```python
#Evaluated at t = 10.
steps = numpy.array([10,100,1000,10000,100000,1000000])
result_analytic = analytic(1000,10)
euler = []
for i in range(len(steps)):
    eulersol = solve_euler(f_rad,1000,0,10,steps[i])
    euler.append(eulersol[])
euler = numpy.array(euler)
euler25 = solve_euler(f_rad,1000,0,6,25)
error_euler = (result_analytic-result_euler)/result_analytic
error_rk4 = (result_analytic-result_rk4)/result_rk4

print(result_analytic)
print(euler)
```

Output:
```
716.5954405062778
[array([1000.        ,  666.75616319,  444.56378115,  296.41564102,
        197.63695551,  131.77565816,   87.86223224,   58.58268486,
         39.06036619,   26.04373989,   17.36482408])
 array([1.00000000e+03, 6.66756163e+02, 4.44563781e+02, 2.96415641e+02,
       1.97636956e+02, 1.31775658e+02, 8.78622322e+01, 5.85826849e+01,
       3.90603662e+01, 2.60437399e+01, 1.73648241e+01, 1.15781035e+01,
       7.71977185e+00, 5.14720546e+00, 3.43193097e+00, 2.28826112e+00,
       1.52571221e+00, 1.01727802e+00, 6.78276387e-01, 4.52244962e-01,
       3.01537115e-01, 2.01051730e-01, 1.34052480e-01, 8.93803174e-02,
       5.95948775e-02, 3.97352519e-02, 2.64937241e-02, 1.76648538e-02,
       1.17781501e-02, 7.85315420e-03, 5.23613897e-03, 3.49122793e-03,
       2.32779774e-03, 1.55207349e-03, 1.03485456e-03, 6.89995658e-04,
       4.60058858e-04, 3.06747079e-04, 2.04525505e-04, 1.36368641e-04,
       9.09246320e-05, 6.06245588e-05, 4.04217982e-05, 2.69514831e-05,
       1.797
```

```python

```

```python

```