# <center>L2 Computational Physics</center>

## <center>Week 3: Differential Equations I</center>

```python
# usual packages to import
import numpy 
import matplotlib.pyplot as plt
%matplotlib inline
```

In this notebook, you will generate and plot the decay curve for Iodine-133 analytically and numerically. $^{133}\textrm{I}$ has a half life $t_{1/2}$ of 20.8 hours. This means that half of the nuclei will have decayed after time $t_{1/2}$. Derive the mean lifetime $\tau$ from that information.

<!-- grade_id: constants, points: 0, grade: false, solution: true -->
```python
# define a function to calculate the mean lifetime from the half life
T_HALF = 20.8
def meanLifetime(halfLife):
    lmb=numpy.log(2)/T_HALF
    TAU=1/lmb
    return TAU
TAU = meanLifetime(T_HALF)

```

Check your average lifetime:

<!-- grade_id: constants_check, points: 1, grade: true, solution: false -->
```python
# this test is worth 1 mark
assert numpy.isclose(TAU, 30.0080568505)         
```

### The Decay Equation

Implement the function `f_rad` such that the differential equation 

$$ \frac{dN}{dt} = f_{rad}(N,t)$$

describes the radioactive decay process.

- *Your function should return values using hours as the time unit.*
- *The function should use the constant* `TAU`.

<!-- grade_id: decay_function, points: 0, grade: false, solution: true -->
```python
def f_rad(N, t):
    # YOUR CODE HERE
    return N*(-1/TAU)


```

Make sure your function works:

<!-- grade_id: decay_test, points: 1, grade: true, solution: false -->
```python
# this test cell is worth 1 mark
assert numpy.isclose(f_rad(1000, 0), -33.324383681)           
```

Solve this first order, ordinary differential equation analytically. Implement this function below, naming it `analytic`. The function should take an initial number of atoms `N0` at time `t=0`, and a time argument. The function should return nuclei count at the time argument. Make sure the function also works for numpy arrays.

<!-- grade_id: analytic_function, points: 0, grade: false, solution: true -->
```python
def analytic(N0, t):
    # YOUR CODE HERE
    N=N0*numpy.exp(-t/TAU)
    return N

```

Check your answer for a single time:

<!-- grade_id: analytic_test, points: 1, grade: true, solution: false -->
```python
# this test is worth 1 mark
assert numpy.isclose(analytic(1000, 41.6), 250.0)           
```

<!-- grade_id: analytic_test_array, points: 2, grade: true, solution: false -->
```python
# this test is worth 1 mark
assert numpy.isclose(analytic(1000, numpy.arange(0, 60, 6)), 
                     [1000.        ,  818.77471839,  670.39203948,  548.90005334,
                       449.4254866 ,  367.97822623,  301.29126855,  246.68967356,
                       201.983268  ,  165.37879338]).all()

```

## Numerically Solving the ODE

We now wish to solve our differential equation numerically. We shall do this using Euler's and RK4 methods.

### Euler's Method

Create a function which takes as its arguments the initial number of atoms, `n0`, the initial time `t0`, the time step, `dt`, and the number of steps to perform, `n_steps`.  This function should return an array of the number of counts at each time step using Euler's method. This array should contain the initial and final values, so the array length should be `n_steps+1` 

<!-- grade_id: euler_function, points: 0, grade: false, solution: true -->
```python

def solve_euler(f, n0, t0, dt, n_panels):
    B=[n0]
    n=n0
    # YOUR CODE HERE
    #x(t)+dt*f(x(t),t)
    for i in range (0,n_panels):
        n=n+dt*f(n,t0+i*dt)
        B.append(n)
    print (B)
    return B

```

Try your solution:

<!-- grade_id: euler_test_length, points: 1, grade: true, solution: false -->
```python
# this test is worth 1 mark
assert len(solve_euler(f_rad, 1000, 0, 1, 17)) == 18
```

Output:
```
[1000, 966.6756163192334, 934.4617471861698, 903.3213853879383, 873.2187569542291, 844.1192800602443, 815.9895252991842, 788.7971772786275, 762.5109974966888, 737.1007884553051, 712.5373589694249, 688.7924906322477, 665.8389053979878, 643.6502342449236, 622.2009868827305, 601.4665224692988, 581.4230213033954, 562.0474574606505]
```

<!-- grade_id: euler_test_one, points: 2, grade: true, solution: false -->
```python
# this test is worth 2 marks
assert numpy.isclose(solve_euler(f_rad, 1000, 0, 6, 1), [1000.,  800.05369792]).all()
```

Output:
```
[1000, 800.0536979154003]
```

<!-- grade_id: euler_test_many, points: 2, grade: true, solution: false -->
```python
# this test is worth 2 mark
assert numpy.isclose(solve_euler(f_rad, 1000, 0, 6, 10), [1000.        ,  800.05369792,  640.08591955,  512.10310692,
                                                409.7099844 ,  327.7899881 ,  262.24959212,  209.81375595,
                                                167.86227132,  134.29883091,  107.4462763 ]).all()
```

Output:
```
[1000, 800.0536979154003, 640.0859195481066, 512.1031069180422, 409.7099844037453, 327.78998809507743, 262.2495921151118, 209.8137559485006, 167.86227132011726, 134.29883091013807, 107.44627629537104]
```

### RK 4 method

Implement the RK4 method in the `solve_RK4` function. The arguments are the same as for `solve_euler`.

<!-- grade_id: RK4, points: 0, grade: false, solution: true -->
```python
def solve_RK4(f, n0, t0, dt, nsteps):
    n=n0
    b=[n0]
    for i in range (1,nsteps+1):
        k1=dt*f( n,t0)
        k2=dt*f(n+0.5*k1,t0+0.5*dt)
        k3=dt*f(n+0.5*k2,t0+0.5*dt) 
        k4=dt*f(n+k3,t0+dt)
        k=( k1+2*k2+2*k3+ k4)/6 
        n=n+k
        t0=t0+dt
        b.append(n)
    return b
  
    

```

<!-- grade_id: RK4-test-1, points: 1, grade: true, solution: false -->
```python
# This checks that we return an array of the right length
# this test is worth 1 mark
assert len(solve_RK4(f_rad, 1000, 0, 1, 17)) == 18
```

<!-- grade_id: RK4-test-2, points: 2, grade: true, solution: false -->
```python
# This checks that a single step is working
# this test is worth 2 mark
assert numpy.isclose(solve_RK4(f_rad, 1000,0, 6, 1), [1000.,  818.7773]).all()
```

<!-- grade_id: RK4-test-3, points: 2, grade: true, solution: false -->
```python
# This checks multiple steps
# this test is worth 2 marks
assert numpy.isclose(solve_RK4(f_rad, 1000, 0, 6, 10), [
    1000.,
    818.77729521,  
    670.39625915,  
    548.90523578,
    449.43114428,  
    367.9840167,  
    301.29695787,  
    246.69510822, 
    201.98835345,  
    165.3834777,  
    135.41223655]).all()
```

## Plotting task

**Task 1: **

Create a plot to show that the RK4 method has an error that scales better with the number of steps than the Euler method. (click on the "+" button to create new cells.)       [task worth 5 marks]


```python
#finding the error in euler
panel_counts =[4, 8, 16, 32, 64, 128, 256, 512, 1024]

euler=[1000, 966.6756163192334, 934.4617471861698, 903.3213853879383, 873.2187569542291, 844.1192800602443, 815.9895252991842, 788.7971772786275, 762.5109974966888, 737.1007884553051, 712.5373589694249, 688.7924906322477, 665.8389053979878, 643.6502342449236, 622.2009868827305, 601.4665224692988, 581.4230213033954, 562.0474574606505, 543.3175723414325, 525.211849100224, 507.7094879271232, 490.7903821530742, 474.4350951513751, 458.62483800892966, 443.3414479415907, 428.5673674287985, 414.2856240435451, 400.47981095449217, 387.1340680778438, 374.23306385732167, 361.76197765131144, 349.70648270694625, 338.0527297015686, 326.78733083266303, 315.89734443798176, 305.3702601281952, 295.1939844149877, 285.35682681808845, 275.8474864352764, 266.65503895993214, 257.7689241312216, 249.17893360249434, 240.87519921396054, 232.84818165617344, 225.08865951129428, 217.5877186595505, 210.33674203871692, 203.32739974485628, 196.55163946294607, 190.00167721639914, 183.6699884248507, 177.549299259939, 171.63257828914954, 165.9130283981227, 160.38407898214572, 155.03937839785831, 149.87278666650053, 144.87836842032038, 140.05038608403817, 135.3832932835342, 130.87172847418793, 126.51050878154899, 122.29462404726367, 118.21923107341755, 114.27964805868177, 110.47134921987129, 106.78995959273634, 103.23125000601443, 99.79113222296887, 96.46565424483255, 93.25099577076158, 90.14346380908317, 87.139488434796, 84.23561868844912, 81.42851861168849, 78.71496341491614, 76.09183577265998, 73.55612224239798, 71.10490980272293, 68.73538250687069, 66.44481824776749, 64.23058563088009, 62.090140951276304, 60.02102527142309, 58.0208615963652, 56.08735214303927, 54.21827570058636, 52.411485078630434, 50.66490464059138, 48.97652791919886, 47.3444153114677, 45.76669185048679, 44.24154505146175, 42.76722282953692, 41.34203148700459, 39.964333767589316, 38.632546975571955, 37.34514115759276, 36.100637345044746, 34.897605855038265, 33.7346646479848, 32.610477739913364, 31.52375366769539, 30.473244005415136, 29.45774193018106, 28.476080835730702, 27.527132992236286, 26.609808250771515, 25.723052790951176, 24.865847910304904, 24.037208853994315, 23.236183683529088, 22.461852183182398, 21.713324802849364, 20.989741636134106, 20.290271432491412, 19.61411064228817, 18.960482493687554, 18.32863610029545, 17.717845598544056, 17.12740931382159, 16.556648954390266, 16.0049088321664, 15.471555109467596, 14.955975070861573, 14.457576419280201, 13.975786595590105, 13.510052120838145, 13.05983796041618, 12.624626909414632, 12.203918998458768, 11.79723091934513, 11.404095469818271, 11.024061016849956, 10.656690977804265, 10.301563318892551, 9.958270070342065, 9.626416857721292, 9.305622448883588, 8.995518316008637, 8.695748212238602, 8.405967762422623, 8.125844067499497, 7.855055322064062, 7.593290444677951, 7.340248720500004, 7.095639455825806, 6.859181644139481, 6.630603643294105, 6.409642863449883, 6.196045465411592, 5.989566069018742, 5.78996747125346, 5.597020373742252, 5.410503119338597, 5.230201437483773, 5.055908198053366, 4.887423173406702, 4.724552808365828, 4.567109997859801, 4.4149138719788565, 4.267789588191494, 4.12556813048582, 3.9880861152043674, 3.855185603349359, 3.726713919142777, 3.6025234746328096, 3.4824716001451774, 3.3664203803845663, 3.254236495997879, 3.145791070417292, 3.0409595218071765, 2.939621419944794, 2.8416603478703535, 2.746963768147501, 2.6554228935805897, 2.566932562240219, 2.4813911166534726, 2.398700287020066, 2.3187650783202445, 2.241493661184738, 2.1667972664014115, 2.0945900829374144, 2.0247891593596794, 1.9573143085405205, 1.8920880155388617, 1.8290353485512645, 1.7680838728304575, 1.7091635674724797, 1.652206744976839, 1.5971479734872802, 1.5439240016238311, 1.492473685819774, 1.442737920080068, 1.3946595680805287, 1.348183397529761, 1.3032560167184397, 1.2598258131830469, 1.2178428944136013, 1.177259030537267, 1.1380275989119957, 1.1001035305665507, 1.063443258425385, 1.028004667258893, 0.9937470453015388, 0.9606310374822822, 0.9286186002135697, 0.8976729576869563, 0.8677585596251476, 0.8388410404419298, 0.8108871797630695, 0.7838648642628302, 0.7577430507722636, 0.732491730616894, 0.7080818951428279, 0.684485502391684, 0.6616754448860612, 0.6396255184885362, 0.618310392298415, 0.5977055795516573, 0.5777874094905429, 0.5585330001707638, 0.5399202321747035, 0.5219277232007051, 0.5045348034991359, 0.48772149212703053, 0.47146847399403335, 0.45575707767327067, 0.44056925395166163, 0.42588755509502735, 0.41169511480417703, 0.3979756288389454, 0.384713336287922, 0.3718930014623555, 0.35949989639343205, 0.34751978391282146, 0.3359389012970535, 0.32474394445693533, 0.3139220526538469, 0.3034607937253563, 0.2933481498031825, 0.2835725035070983, 0.27412262459891223, 0.2649876570811993, 0.25615710672595804, 0.24762082901886714, 0.2393690175052929, 0.2313921925246584, 0.22368119032023287, 0.21622715251183086, 0.20902151591932697, 0.20205600272529586, 0.1953226109654761, 0.18881360533613345, 0.18252150830776329, 0.17643909153492315, 0.17055936755232748, 0.1648755817476648, 0.15938120460191604, 0.15406992418825902, 0.14893563892094286, 0.14397245054580124, 0.13917465736435275, 0.13453674768370383, 0.1300533934847296, 0.12571944430125875, 0.12152992130323084, 0.1174800115770286, 0.1135650625964148, 0.1097805768777216, 0.1061222068131525, 0.10258574967626134, 0.09916714279387052, 0.09586245887888221, 0.09266790151862063, 0.08957980081352263, 0.08659460916116615, 0.08370889718079343, 0.08091934977364683, 0.07822276231459167, 0.07561603697065081, 0.0730961791422218, 0.07066029402288836, 0.06830558327387384, 0.06602934180931672, 0.06382895468867456, 0.061701894112686895, 0.05964571651944568, 0.05765805977723744, 0.0557366404709322, 0.053879251278801915, 0.05208375843675469, 0.05034809928707191, 0.048670279908832194, 0.04704837282729996, 0.045480514799647265, 0.04396490467446504, 0.042499801322604835, 0.041083521636974, 0.0397144385989864, 0.03839097940944753, 0.03711162368172669, 0.035874901695140606, 0.03467939270654195, 0.033523723318173174, 0.032406565899910506, 0.03132663706408584, 0.030282696191134117, 0.029273544004372674, 0.028298021192275154, 0.027355007076657313, 0.0264434183252447, 0.02556220770714323, 0.024710362889782938, 0.023886905275952833, 0.023090888879590852, 0.02232139923903742, 0.021577552366504167, 0.02085849373255095, 0.020163397284404555, 0.01949146449699133, 0.01884192345559355, 0.018214027969075713, 0.01760705671266202, 0.017020312399280254, 0.01645312097852013, 0.015904830862285856, 0.015374812176253343, 0.014862456036272155, 0.014367173848880897, 0.013888396635132513, 0.01342557437695269, 0.01297817538528045, 0.012545685689265083, 0.01212760844581771, 0.011723463368839176, 0.011332786177468567, 0.010955128062718516, 0.010590055171884551, 0.010237148110136183, 0.00989600145871717, 0.009566223309201453, 0.009247434813269732, 0.008939269747489452, 0.008641374092598244, 0.008353405626807463, 0.008075033532658657, 0.007805938016981284, 0.007545809943515117, 0.007294350477775276, 0.007051270743751909, 0.006816291492050156, 0.006589142779089132, 0.006369563656991413, 0.006157301873806764, 0.005952113583725724, 0.005753763066950145, 0.005562022458898874, 0.005376671488437487, 0.005197497224831358, 0.005024293833131358, 0.004856862337711179, 0.004695010393684627, 0.0045385520659402935, 0.004387307615539763, 0.004241103293233967, 0.004099771139860476, 0.003963148793392431, 0.0038310793024174546, 0.003703410945832252, 0.0035799970585457873, 0.0034606958629907914, 0.0033453703062500447, 0.0032338879026103245, 0.0031261205813631486, 0.003021944539677462, 0.0029212401003752527, 0.0028238915744467066, 0.0027297871281469605, 0.0026388186545217734, 0.0025508816492145255, 0.002465875090411874, 0.002383701322790144, 0.002304265945329134, 0.002227477702864462, 0.002153248381253854, 0.002081492706036961, 0.002012128244472268, 0.001945075310838567, 0.0018802568748921963, 0.0018175984733748895, 0.0017570281244705692, 0.0016984762451128141, 0.001641875571048007, 0.0015871610795623254, 0.0015342699147838108, 0.001483141315473698, 0.0014337165452240556, 0.001385938824981546, 0.0013397532678197902, 0.0012951068158854027, 0.0012519481794452615, 0.0012102277779649903, 0.0011698976831509634, 0.0011309115638904008, 0.0010932246330263014, 0.0010567935959060676, 0.0010215766006447169, 0.000987533190045539, 0.0009546242551229701, 0.0009228119901742863, 0.0008920598493485065, 0.00086233250466261, 0.0008335958054168368, 0.0008058167389624484, 0.0007789633927768796, 0.0007530049178027112, 0.0007279114930083495, 0.0007036542911296995, 0.0006802054455534755, 0.0006575380183041047, 0.0006356259690974478, 0.0006144441254257853, 0.0005939681536397033, 0.0005741745309936574, 0.0005550405186231005, 0.0005365441354221326, 0.0005186641327916603, 0.0005013799702290589, 0.0004846717917312944, 0.0004685204029843961, 0.0004529072493130767, 0.0004378143943651671, 0.0004232244995063798, 0.0004091208039017288, 0.0003954871052607239, 0.0003823077412242198, 0.0003695675713715366, 0.00035725195982718247, 0.00034534675844719563, 0.0003338382905657923, 0.0003227133352836466, 0.00031195911227975445, 0.0003015632671294326, 0.0002915138571115859, 0.00028179933748893925, 0.000272408548245472, 0.0002633307012658193, 0.0002545553679419118, 0.0002460724671926168, 0.0002378722538826172, 0.00022994530762722413, 0.0002222825219702626, 0.00021487509392259714, 0.00020771451384927976, 0.00020079255569370246, 0.00019410126752752383, 0.00018763296241551352, 0.0001813802095848201, 0.0001753358258885177, 0.00016949286755362463, 0.00016384462220411428, 0.00015838460114975413, 0.00015310653193191453, 0.00014800435111778386, 0.00014307219733471194, 0.00013830440453667964, 0.00013369549549515936, 0.00012924017550688846, 0.0001249333263113273, 0.00012077000021081422, 0.00011674541438666278, 0.00011285494540467154, 0.0001090941239037343, 0.00010545862946144917, 0.00010194428563082804, 9.854705514240467e-05, 9.526303526622951e-05, 9.208845332842328e-05, 8.901966237713854e-05, 8.605313699295047e-05, 8.318546923886381e-05, 8.041336474528331e-05, 7.773363892545006e-05, 7.514321331699618e-05, 7.263911204541492e-05, 7.021845840538331e-05, 6.787847155601036e-05, 6.561646332621387e-05, 6.342983512655616e-05, 6.131607496399104e-05, 5.927275455609235e-05, 5.729752654144922e-05, 5.538812178302306e-05, 5.354234676136857e-05, 5.175808105472407e-05, 5.003327490307623e-05, 4.836594685340084e-05, 4.6754181483374545e-05, 4.519612720094238e-05, 4.368999411721344e-05, 4.2234051990240986e-05, 4.082662823732475e-05, 3.946610601355212e-05, 3.81509223543707e-05, 3.687956638005852e-05, 3.5650577560029146e-05, 3.44625440349778e-05, 3.3314100994940885e-05, 3.220392911140567e-05, 3.1130753021668975e-05, 3.0093339863703693e-05, 2.9090497859849922e-05, 2.812107494770376e-05, 2.7183957456630886e-05, 2.627806882838448e-05, 2.5402368380357802e-05, 2.4555850110050586e-05, 2.3737541539375866e-05, 2.294650259747957e-05, 2.2181824540789453e-05, 2.144262890905274e-05, 2.0728066516163168e-05, 2.0037316474618096e-05, 1.9369585252484977e-05, 1.872410576179385e-05, 1.810013647730858e-05, 1.7496960584664508e-05, 1.6913885156893897e-05, 1.6350240358393143e-05, 1.5805378675417295e-05, 1.5278674172217884e-05, 1.4769521771969476e-05, 1.427733656165893e-05, 1.380155312013877e-05, 1.3341624868572785e-05, 1.2897023442527608e-05, 1.2467238084988977e-05, 1.205177505960534e-05, 1.1650157083484758e-05, 1.126192277889351e-05, 1.0886626143226498e-05, 1.0523836036640554e-05, 1.0173135686762066e-05, 9.834122209899907e-06, 9.506406148213655e-06, 9.189611022305384e-06, 8.883372898721077e-06, 8.587339971864772e-06, 8.301172159845166e-06, 8.024540713790387e-06, 7.757127840182104e-06, 7.498626335775119e-06, 7.248739234683048e-06, 7.007179467224643e-06, 6.7736695301388595e-06, 6.547941167789794e-06, 6.32973506399528e-06, 6.1188005441251e-06, 5.914895287126592e-06, 5.7177850471468274e-06, 5.5272433844315565e-06, 5.3430514051917805e-06, 5.16499751013911e-06, 4.99287715140103e-06, 4.826492597536809e-06, 4.665652706384113e-06, 4.510172705475362e-06, 4.359873979771579e-06, 4.2145838664698805e-06, 4.074135456648869e-06, 3.938367403524087e-06, 3.8071237370932257e-06, 3.680253684958177e-06, 3.5576114991180755e-06, 3.4390562885343576e-06, 3.3244518572754855e-06, 3.2136665480554002e-06, 3.1065730909859574e-06, 3.0030484573695963e-06, 2.9029737183642777e-06, 2.8062339083583247e-06, 2.7127178928982145e-06, 2.6223182410175936e-06, 2.5349311018208504e-06, 2.450456085179464e-06, 2.3687961464040744e-06, 2.2898574747597835e-06, 2.2135493856966172e-06, 2.1397842166713377e-06, 2.0684772264409335e-06, 1.999546497712088e-06, 1.9329128430347972e-06, 1.8684997138320243e-06, 1.8062331124608833e-06, 1.7460415072043317e-06, 1.6878557500957105e-06, 1.631608997481733e-06, 1.5772366332326608e-06, 1.524676194511455e-06, 1.4738673000166242e-06, 1.4247515806163346e-06, 1.3772726122940971e-06, 1.331375851328997e-06, 1.287008571636002e-06, 1.2441198041943686e-06, 1.2026602784945553e-06, 1.1625823659363852e-06, 1.1238400251134276e-06, 1.0863887489207453e-06, 1.0501855134252423e-06, 1.0151887284398767e-06, 9.813581897449567e-07, 9.486550329016332e-07, 9.170416886045289e-07, 8.864818395222135e-07, 8.569403785759435e-07, 8.283833686087374e-07, 8.007780033984539e-07, 7.740925699700856e-07, 7.482964121639718e-07, 7.233598954180786e-07, 6.992543727238873e-07, 6.759521517167827e-07, 6.534264628631328e-07, 6.316514287075156e-07, 6.10602034144762e-07, 5.902540976826654e-07, 5.705842436623436e-07, 5.515698754043396e-07, 5.331891492496127e-07, 5.154209494655971e-07, 4.982448639885005e-07, 4.816411609739764e-07, 4.6559076612922975e-07, 4.5007524080051724e-07, 4.350767607908674e-07, 4.2057809588368745e-07, 4.0656259004873323e-07, 3.93014142307703e-07, 3.7991718823747373e-07, 3.672566820897301e-07, 3.5501807950644664e-07, 3.4318732081136493e-07, 3.317508148582727e-07, 3.206954234175286e-07, 3.10008446082897e-07, 2.996776056813523e-07, 2.8969103416909344e-07, 2.800372589975645e-07, 2.7070518993381945e-07, 2.6168410632009006e-07, 2.5296364475792084e-07, 2.4453378720272273e-07, 2.3638484945506826e-07, 2.285074700355073e-07, 2.208925994301228e-07, 2.135314896944715e-07, 2.0641568440396728e-07, 1.9953700893916144e-07, 1.9288756109476027e-07, 1.8645970200159117e-07, 1.8024604735108875e-07, 1.7423945891221943e-07, 1.6843303633109947e-07, 1.628201092038854e-07, 1.573942294138308e-07, 1.5214916372370572e-07, 1.470788866150692e-07, 1.4217757336616867e-07, 1.374395933605141e-07, 1.3285950361843978e-07, 1.284320425442227e-07, 1.2415212388157448e-07, 1.2001483087056283e-07, 1.1601541059924988e-07, 1.121492685435588e-07, 1.0841196328909592e-07, 1.047992014288649e-07, 1.0130683263101147e-07, 9.793084487093244e-08, 9.466735982227186e-08, 9.151262840150928e-08, 8.846302646102197e-08, 8.551505062567307e-08, 8.266531426814296e-08, 7.991054361838022e-08, 7.724757400270268e-08, 7.467334620822821e-08, 7.218490296845849e-08, 6.977938556597867e-08, 6.745403054836985e-08, 6.520616655356183e-08, 6.303321124097896e-08, 6.093266832495377e-08, 5.890212470700012e-08, 5.6939247703651684e-08, 5.504178236668099e-08, 5.3207548892620456e-08, 5.1434440118609624e-08, 4.972041910169166e-08, 4.8063516778778375e-08, 4.64618297045954e-08, 4.491351786500903e-08, 4.34168025632225e-08, 4.1969964376413587e-08, 4.0571341180465874e-08, 3.9219326240524746e-08, 3.7912366365184345e-08, 3.664896012218515e-08, 3.5427656113572336e-08, 3.4247051308333395e-08, 3.3105789430599596e-08, 3.2002559401559626e-08, 3.093609383329553e-08, 2.990516757281059e-08, 2.890859629457663e-08, 2.794523513998377e-08, 2.701397740212971e-08, 2.6113753254437582e-08, 2.5243528521641835e-08, 2.4402303491730267e-08, 2.3589111767477336e-08, 2.2803019156249433e-08, 2.2043122596806706e-08, 2.1308549121868544e-08, 2.0598454855250933e-08, 1.99120240424236e-08, 1.9248468113373228e-08, 1.8607024776696175e-08, 1.798695714388002e-08, 1.7387552882767857e-08, 1.6808123399232883e-08, 1.6248003046123176e-08, 1.5706548358567904e-08, 1.518313731476647e-08, 1.4677168621411428e-08, 1.4188061022924206e-08, 1.371525263371015e-08, 1.3258200292665748e-08, 1.2816378939196503e-08, 1.2389281010028621e-08, 1.1976415856121592e-08, 1.1577309179011779e-08, 1.119150248593953e-08, 1.0818552563133828e-08, 1.0458030966649414e-08, 1.010952353017145e-08, 9.772629889222279e-09, 9.446963021223707e-09, 9.132148800886435e-09, 8.827825570415843e-09, 8.533643724040422e-09, 8.249265306385533e-09, 7.974363624231106e-09, 7.708622871207279e-09, 7.451737764996835e-09, 7.203413196627622e-09, 6.963363891452106e-09, 6.73131408142456e-09, 6.506997188299421e-09, 6.290155517386861e-09, 6.080539961513771e-09, 5.877909714850052e-09, 5.682031996271483e-09, 5.49268178194134e-09, 5.30964154680357e-09, 5.132701014690548e-09, 4.961656916758341e-09, 4.7963127579719565e-09, 4.636478591372343e-09, 4.481970799865791e-09, 4.332611885285071e-09, 4.188230264479982e-09, 4.0486600722030525e-09, 3.913740970563958e-09, 3.783317964833748e-09, 3.6572412253872914e-09, 3.535365915579368e-09, 3.4175520253566967e-09, 3.303664210414729e-09, 3.193571636714452e-09, 3.0871478301805658e-09, 2.9842705314083827e-09, 2.8848215552125246e-09, 2.7886866548560767e-09, 2.695755390804219e-09, 2.6059210038515642e-09, 2.5190802924774463e-09, 2.4351334942882702e-09, 2.353984171410722e-09, 2.27553909970418e-09, 2.1997081616650514e-09, 2.1264042429000116e-09, 2.0555431320492016e-09, 1.9870434240444294e-09, 1.9208264265912287e-09, 1.8568160697673466e-09, 1.7949388186338065e-09, 1.7351235887581516e-09, 1.6773016645528262e-09, 1.6214066203348794e-09, 1.567374244016305e-09, 1.515142463337354e-09, 1.4646512745580782e-09, 1.415842673526181e-09, 1.3686605890419921e-09, 1.3230508184440127e-09, 1.2789609653410323e-09, 1.236340379419284e-09, 1.1951400982554913e-09, 1.1553127910689562e-09, 1.116812704348077e-09, 1.079595609288827e-09, 1.0436187509848152e-09, 1.0088407993105547e-09, 9.752218014415186e-10, 9.42723135956433e-10, 9.113074684690853e-10, 8.809387087386733e-10, 8.515819692094266e-10, 8.23203524931869e-10, 7.957707748196799e-10, 7.692522041976479e-10, 7.436173485976901e-10, 7.188367587613463e-10, 6.948819668085445e-10, 6.717254535337709e-10, 6.493406167920746e-10, 6.277017409385899e-10, 6.067839672864671e-10, 5.865632655492752e-10, 5.670164062350677e-10, 5.481209339604008e-10, 5.298551416536443e-10, 5.121980456179513e-10, 4.951293614252399e-10, 4.786294806134922e-10, 4.6267944816060217e-10, 4.472609407088929e-10, 4.323562455152892e-10, 4.1794824010296197e-10, 4.040203725910697e-10, 3.905566426799986e-10, 3.7754158327025827e-10, 3.649602426939161e-10, 3.527981675381583e-10, 3.4104138604124535e-10, 3.2967639204178645e-10, 3.186901294628951e-10, 3.080699773134004e-10, 2.978037351888836e-10, 2.878796092558838e-10, 2.782861987031716e-10, 2.6901248264452503e-10, 2.600478074579633e-10, 2.51381874546892e-10, 2.4300472850910105e-10, 2.3490674570002325e-10, 2.2707862317711541e-10, 2.19511368012661e-10, 2.1219628696271712e-10, 2.051249764803375e-10, 1.982893130615985e-10, 1.9168144391333813e-10, 1.852937779318867e-10, 1.7911897698242574e-10, 1.73149947468957e-10, 1.6737983218519687e-10, 1.6180200243703504e-10, 1.5641005042750694e-10, 1.5119778189553265e-10, 1.4615920899996504e-10, 1.4128854344077286e-10, 1.365801898094559e-10, 1.3202873916105365e-10, 1.2762896280036283e-10, 1.2337580627522526e-10, 1.1926438356998572e-10, 1.152899714924494e-10, 1.1144800424789038e-10, 1.0773406819387797e-10, 1.0414389676989531e-10, 1.0067336559592517e-10, 9.731848773437247e-11, 9.407540910988026e-11, 9.094040408177753e-11, 8.790987116407242e-11, 8.49803288880741e-11, 8.214841180289018e-11, 7.941086660920506e-11, 7.676454842189774e-11, 7.420641715720562e-11, 7.173353404028388e-11, 6.934305822914813e-11, 6.703224355112226e-11, 6.479843534804207e-11, 6.263906742659056e-11, 6.055165911026145e-11, 5.853381238956412e-11, 5.658320916719627e-11, 5.469760859501955e-11, 5.2874844499778725e-11, 5.1112822894607225e-11, 4.940951957346026e-11, 4.776297778571192e-11, 4.617130598824492e-11, 4.463267567245057e-11, 4.314531926364261e-11, 4.170752809047181e-11, 4.031765042200858e-11, 3.8974089570238543e-11, 3.767530205579135e-11, 3.6419795834795385e-11, 3.5206128584821476e-11, 3.403290604794648e-11, 3.289878042903323e-11, 3.180244884738683e-11, 3.074265184000856e-11, 2.971817191472789e-11, 2.8727832151550516e-11, 2.7770494850615585e-11, 2.684506022520892e-11, 2.595046513833077e-11, 2.5085681881366677e-11, 2.424971699345836e-11, 2.3441610120218347e-11, 2.2660432910477248e-11, 2.1905287949796233e-11, 2.117530772951955e-11, 2.0469653650182736e-11, 1.9787515058131643e-11, 1.9128108314245517e-11, 1.849067589369434e-11, 1.7874485515696166e-11, 1.7278829302274802e-11, 1.6703022965051323e-11, 1.6146405019135298e-11, 1.560833602321258e-11, 1.5088197844956714e-11, 1.458539295092006e-11, 1.4099343720088851e-11, 1.3629491780313604e-11, 1.3175297366852578e-11, 1.2736238702291388e-11, 1.2311811397126401e-11, 1.1901527870323326e-11, 1.1504916789185334e-11, 1.1121522527887229e-11, 1.0750904644053625e-11, 1.0392637372779846e-11, 1.0046309137514257e-11, 9.711522077240141e-12, 9.387891589413955e-12, 9.075045888134882e-12, 8.772625577038112e-12, 8.480283236421187e-12, 8.197683024129114e-12, 7.924500289739728e-12, 7.660421201606096e-12, 7.405142386327495e-12, 7.15837058023481e-12, 6.919822292489953e-12, 6.689223479412296e-12, 6.466309229657968e-12, 6.250823459890363e-12, 6.04251862059224e-12, 5.841155411681447e-12, 5.646502507603588e-12, 5.458336291585795e-12, 5.276440598746337e-12, 5.10060646776494e-12, 4.9306319008285416e-12, 4.766321631576703e-12, 4.607486900780104e-12, 4.453945239494402e-12, 4.3055202594403665e-12, 4.162041450369462e-12, 4.023343984182096e-12, 3.8892685255735076e-12, 3.759661048989767e-12, 3.6343726616835985e-12, 3.513259432666765e-12, 3.3961822273625052e-12, 3.2830065477680767e-12, 3.1736023779437842e-12, 3.0678440346509924e-12, 2.9656100229675315e-12, 2.8667828967146346e-12, 2.7712491235350565e-12, 2.678898954467386e-12, 2.5896262978667105e-12, 2.503328597526797e-12, 2.4199067148637785e-12, 2.3392648150259945e-12, 2.2613102567991507e-12, 2.185953486180323e-12, 2.1131079334985405e-12, 2.042689913963763e-12, 1.9746185315300027e-12, 1.908815585962145e-12, 1.845205482999715e-12, 1.7837151475143784e-12, 1.724273939561414e-12, 1.6668135732287225e-12, 1.611268038190139e-12, 1.5575735238729347e-12, 1.5056683461523894e-12, 1.455492876489222e-12, 1.4069894734284725e-12, 1.3601024163811423e-12, 1.3147778416125194e-12, 1.2709636803636535e-12, 1.228609599034896e-12, 1.1876669413627844e-12, 1.1480886725238484e-12, 1.1098293251011218e-12, 1.0728449468512856e-12, 1.0370930502124418e-12, 1.0025325634945058e-12, 9.691237836961524e-13, 9.368283308941055e-13, 9.056091041523782e-13, 8.75430238900809e-13, 8.462570657339333e-13]
E=numpy.array(euler)

RK4=[1000, 967.2247571093798, 935.5237307652988, 904.8617132595269, 875.2046508250232, 846.5196058152326, 818.7747201229663, 791.9391797982365, 765.9831808257508, 740.8778960240569, 716.5954430295767, 693.1088533299708, 670.3920423124417, 648.4197802937125, 627.1676644995035, 606.6120919623893, 586.7302333079347, 567.5000073999969, 548.9000568170333, 530.9097241321798, 513.5090289707555, 496.6786458197124, 480.399882564387, 464.6546597287138, 449.42549039584674, 434.69546078688677, 420.44821147614647, 406.6679192220889, 393.33927939376184, 380.4474889732098, 367.9782301149863, 355.91765424450705, 344.25236667758355, 332.9694117440549, 322.0562583989966, 311.50078630552514, 301.29127237374234, 291.41637774086894, 281.8651351781072, 272.62693691024725, 263.6915228344881, 255.04896912539024, 246.68967721330327, 238.60436312402857, 230.7840471678768, 223.22004396666927, 215.90395280760674, 208.82764831329243, 201.98327141754726, 195.36322063699507, 188.96014362872373, 182.76692902464583, 176.77669853349033, 170.98279930165324, 165.3787965244234, 159.95846629937697, 154.7157887140038, 149.64494115988845, 144.74029186602053, 139.99639364405243, 135.40797783855774, 130.9699484755713, 126.67737660291243, 122.5254948160054, 118.5096919631174, 114.62550802413365, 110.86862915718193, 107.2348829076052, 103.72023357396122, 100.32077772590279, 97.0327398689604, 93.85246825141286, 90.77643080858857, 87.8012112400935, 84.92350521560877, 82.14011670504433, 79.44795442897262, 76.84402842540011, 74.3254467290639, 71.88941215956497, 69.53321921477132, 67.25425106604045, 65.04997665192423, 62.917947867128234, 60.85579684360373, 58.861233320752376, 56.93204210182325, 55.06608059367697, 53.26127642718474, 51.5156251556193, 49.82718802848173, 48.19408983829164, 46.61451683794926, 45.08671472635657, 43.60898670006013, 42.17969156875183, 40.79724193253454, 39.46010241892832, 38.1667879776592, 36.91586223133661, 35.70593588018788, 34.53566515907781, 33.4037503450999, 32.30893431408161, 31.250001144400493, 30.225774766560605, 29.235117657029402, 28.276929574884402, 27.350146339866605, 26.45373865048347, 25.586710940848885, 24.748100274990474, 23.936975277396236, 23.152435098612802, 22.393608414746446, 21.659652459755694, 20.94975208946078, 20.263118876230426, 19.598990233340462, 18.956628568031835, 18.33532046232732, 17.734375880697183, 17.153127403693777, 16.59092948670396, 16.047157742996085, 15.52120825026529, 15.012496879906948, 14.52045864827332, 14.044547089192955, 13.584233647055903, 13.13900708979071, 12.708372941081238, 12.291852931192714, 11.88898446579709, 11.49932011220778, 11.122427102453177, 10.757886852637057, 10.405294498052067, 10.064258443529976, 9.734399928529305, 9.41535260647732, 9.106762137899192, 8.808285796882442, 8.51959209043962, 8.240360390346455, 7.970280577046603, 7.709052695227508, 7.456386620684836, 7.21200173810552, 6.975626629411536, 6.746998772318295, 6.525864248772846, 6.311977462948101, 6.105100868479856, 5.905004704643693, 5.71146674117874, 5.524272031474908, 5.343212673849457, 5.168087580647801, 4.998702254912072, 4.8348685743694375, 4.676404582500252, 4.523134286453997, 4.3748874615885756, 4.231499462415881, 4.092811039743672, 3.958668163810661, 3.828921853218401, 3.703428009469964, 3.5820472569316597, 3.4646447880400446, 3.3510902135823106, 3.24125741788377, 3.135024418741605, 3.0322732319493233, 2.932889740261458, 2.836763566652981, 2.743787951732667, 2.6538596351742716, 2.566878741033822, 2.482748666825669, 2.401375976234094, 2.322670295341321, 2.2465442122566808, 2.172913180035451, 2.101695422779559, 2.032811844815854, 1.9661859428510844, 1.9017437210060169, 1.8394136086343327, 1.7791263808350302, 1.720815081570052, 1.664414949301751, 1.6098633450676068, 1.5570996829123096, 1.506065362599951, 1.4567037045315876, 1.4089598867958986, 1.3627808842830222, 1.318115409793952, 1.2749138570800858, 1.2331282457496684, 1.1927121679799386, 1.1536207369757978, 1.1158105371177598, 1.0792395757438118, 1.0438672365116384, 1.0096542342894088, 0.9765625715250302, 0.9445554960454087, 0.9135974602388501, 0.883654081575268, 0.8546921044203507, 0.8266793631012783, 0.7995847461829707, 0.7733781619151889, 0.7480305048121172, 0.7235136233273067, 0.6998002885880815, 0.676864164154681, 0.6546797767705547, 0.6332224880713228, 0.6124684672209822, 0.5923946644449687, 0.5729787854306774, 0.5541992665670143, 0.5360352509954769, 0.5184665654461655, 0.5014736978330018, 0.48503777558326766, 0.4691405446773999, 0.45376434937576027, 0.4388921126098655, 0.42450731701629973, 0.410593986592245, 0.3971366689522561, 0.38412041816657405, 0.371530778161918, 0.35935376666631996, 0.3475758596801721, 0.3361839764562383, 0.32516546497195054, 0.3145080878778534, 0.3042000089065922, 0.2942297797273498, 0.28458632723113225, 0.2752589412327824, 0.266237262576063, 0.2575112716285987, 0.2490712771538989, 0.24090790554810287, 0.2330120904294932, 0.22537506256921538, 0.21798834015202062, 0.21084371935621501, 0.2039332652423533, 0.1972493029405579, 0.19078440912667557, 0.18453140377780533, 0.17848334219804066, 0.1726335073055702, 0.16697540217257048, 0.1615027428096055, 0.1562094511865193, 0.15108964848207065, 0.14613764855481237, 0.1413479516279643, 0.13671523818126613, 0.13223436304302613, 0.1279003496758045, 0.12370838464938474, 0.11965381229489488, 0.11573212953414103, 0.11193898087841084, 0.10827015359119244, 0.10472157300943635, 0.10128929801816426, 0.09796951667339851, 0.09475854196855121, 0.09165280773957092, 0.08864886470429917, 0.08574337663163803, 0.08293311663627416, 0.08021496359484413, 0.07758589867956085, 0.07504300200545119, 0.07258344938748124, 0.0702045092039675, 0.06790353936279067, 0.06567798436704242, 0.06352537247684624, 0.06144331296420048, 0.05942949345779442, 0.05748167737484868, 0.05559770143712775, 0.0537754732683657, 0.05201296907043696, 0.05030823137569107, 0.04865936687295527, 0.047064544304790366, 0.045521992433664504, 0.04402999807478617, 0.04258690419341151, 0.04119110806451288, 0.03984105949276468, 0.03853525909086967, 0.03727225661431344, 0.03605064935069779, 0.03486908056186409, 0.033726237977076384, 0.032620852335590846, 0.0315516959769928, 0.030517581477735853, 0.029517360332368767, 0.028549921677985424, 0.02761419106048127, 0.026709129241246, 0.025833731042967194, 0.02498702423326299, 0.024168068444903983, 0.02337595413142512, 0.022609801556967665, 0.021868759819229328, 0.02115200590443745, 0.02045874377329568, 0.01978820347688895, 0.019139640301564897, 0.018512333941842003, 0.017905587700425857, 0.017318727714445096, 0.016751102207047643, 0.016202080763526052, 0.01567105363116804, 0.01515743104205457, 0.014660642558053405, 0.014180136437280642, 0.013715379021326635, 0.013265854142565738, 0.012831062550891606, 0.012410521359241393, 0.012003763507293026, 0.011610337242739934, 0.011229805619567119, 0.010861746012771355, 0.010505749648986549, 0.010161421152492967, 0.009828378106106123, 0.00950625062645764, 0.00919468095319638, 0.008893323051653608, 0.00860184222853091, 0.008319914760184016, 0.00804722753308973, 0.007783477696096627, 0.007528372324073335, 0.0072816280925808085, 0.007042970963207309, 0.006812135879216604, 0.006588866471171371, 0.0063729147722048655, 0.006164040942624629, 0.0059620130035423794, 0.005766606579234242, 0.005577604647945191, 0.005394797300860935, 0.0052179815089795556, 0.005046960897623985, 0.004881545528344897, 0.004721551687971772, 0.004566801684577879, 0.004417123650132546, 0.004272351349621549, 0.004132323996423634, 0.003996886073738111, 0.003865887161865207, 0.0037391817711473446, 0.003616629180385811, 0.0034980932805533612, 0.003383442423629178, 0.003272549276388303, 0.003165290678983153, 0.003061547508160064, 0.0029612045449589446, 0.0028641503467491066, 0.0027702771234591507, 0.0026794806178634485, 0.0025916599897922647, 0.002506717704136921, 0.0024245594225256154, 0.0023450938985495966, 0.002268232876423322, 0.0021938909929660575, 0.002121985682796051, 0.0020524370866319918, 0.001985167962599911, 0.0019201036004470214, 0.0018571717385672158, 0.00179630248374608, 0.0017374282335362778, 0.001680483601177105, 0.0016254053429748212, 0.0015721322880631096, 0.0015206052704656546, 0.0014707670633853855, 0.0014225623156474052, 0.0013759374902250184, 0.0013308408047805831, 0.001287222174155151, 0.0012450331547430237, 0.0012042268906894459, 0.001164758061851683, 0.001126582833465686, 0.0010896588074624449, 0.0010539449753799598, 0.0010194016728185327, 0.0009859905353888007, 0.0009536744561035801, 0.0009224175441662051, 0.0008921850851095883, 0.0008629435022417329, 0.0008346603193548777, 0.000807304124656859, 0.0007808445358846309, 0.0007552521665611985, 0.0007304985933584881, 0.0007065563245299073, 0.0006833987693775357, 0.0006610002087200361, 0.0006393357663284862, 0.0006183813812984093, 0.0005981137813273167, 0.0005785104568680866, 0.0005595496361294714, 0.0005412102608959697, 0.0005234719631402084, 0.0005063150424018583, 0.0004897204439079627, 0.00047366973741037687, 0.00045814509671681546, 0.00044312927989277514, 0.0004286056101123438, 0.00041455795713662925, 0.0004009707193992369, 0.0003878288066789002, 0.00037511762334001985, 0.0003628230521224985, 0.00035093143846286743, 0.0003394295753292922, 0.0003283046885536146, 0.00031754442264414045, 0.000307136827063417, 0.0002970703429557591, 0.0002873337903097842, 0.0002779163555416985, 0.0002688075794855434, 0.000259997345777065, 0.00025147586961830513, 0.00024323368691043522, 0.00023526164374276464, 0.00022755088622624896, 0.00022009285066020777, 0.00021287925402133047, 0.00020590208476440732, 0.0001991535939245688, 0.00019262628651115112, 0.00018631291318362992, 0.00018020646220037742, 0.00017430015163130068, 0.00016858742182571286, 0.00016306192812707166, 0.00015771753382649403, 0.00015254830334722108, 0.0001475484956524639, 0.00014271255786930877, 0.00013803511912160047, 0.00013351098456495432, 0.0001291351296172721, 0.00012490269437835427, 0.0001208089782324108, 0.00011684943462747589, 0.00011301966602592871, 0.00010931541902051213, 0.00010573257961042492, 0.0001022671686322414, 9.891533734058367e-05, 9.567336313363841e-05, 9.25376454187709e-05, 8.95047016136446e-05, 8.657116327840491e-05, 8.373377237463165e-05, 8.09893776469052e-05, 7.833493112296771e-05, 7.576748472859244e-05, 7.328418701340147e-05, 7.08822799839956e-05, 6.85590960408792e-05, 6.631205501577802e-05, 6.413866130605973e-05, 6.20365011030744e-05, 6.000323971133691e-05, 5.803661895557373e-05, 5.6134454672754424e-05, 5.4294634286322384e-05, 5.251511445993077e-05, 5.079391882807782e-05, 4.912913580112112e-05, 4.751891644223311e-05, 4.596147241393983e-05, 4.445507399196241e-05, 4.2998048144155353e-05, 4.158877667240808e-05, 4.022569441544614e-05, 3.890728751053603e-05, 3.763209171216301e-05, 3.639869076581477e-05, 3.520571483506462e-05, 3.405183898020747e-05, 3.2935781686758876e-05, 3.1856303442182915e-05, 3.0812205359268066e-05, 2.9802327844622384e-05, 2.882554931080899e-05, 2.7880784930691676e-05, 2.6966985432607113e-05, 2.6083135935025596e-05, 2.5228254819406068e-05, 2.4401392639993573e-05, 2.360163106934839e-05, 2.2828081878435688e-05, 2.2079885950142994e-05, 2.1356212325129863e-05, 2.0656257278950074e-05, 1.9979243429421345e-05, 1.9324418873251233e-05, 1.8691056350960337e-05, 1.8078452439175343e-05, 1.7485926769394845e-05, 1.691282127236033e-05, 1.635849944719307e-05, 1.582234565448524e-05, 1.5303764432560138e-05, 1.4802179836142144e-05, 1.4317034796701945e-05, 1.3847790503766576e-05, 1.3393925806507202e-05, 1.2954936634939982e-05, 1.253033544009723e-05, 1.2119650652547096e-05, 1.1722426158660402e-05, 1.1338220794042946e-05, 1.0966607853570709e-05, 1.0607174617483745e-05, 1.0259521893012494e-05, 9.923263571027374e-06, 9.598026197219309e-06, 9.28344855733491e-06, 8.97918127600568e-06, 8.684886428725686e-06, 8.400237166546751e-06, 8.124917353074367e-06, 7.858621213361139e-06, 7.601052994307847e-06, 7.3519266361949314e-06, 7.110965454979622e-06, 6.877901835005855e-06, 6.652476931785696e-06, 6.434440384522172e-06, 6.223550038054242e-06, 6.019571673915085e-06, 5.822278750205021e-06, 5.6314521499901546e-06, 5.4468799379473215e-06, 5.268357124985052e-06, 5.0956854405791375e-06, 4.928673112569959e-06, 4.76713465417701e-06, 4.610890657994066e-06, 4.459767596736219e-06, 4.313597630517472e-06, 4.172218420444857e-06, 4.0354729483220575e-06, 3.903209342262275e-06, 3.7752807080166908e-06, 3.651544965831171e-06, 3.531864692650033e-06, 3.416106969491622e-06, 3.3041432338261937e-06, 3.195849136792141e-06, 3.0911044050919993e-06, 2.989792707414843e-06, 2.8918015252367164e-06, 2.797022027855617e-06, 2.705348951522234e-06, 2.6166804825322077e-06, 2.530918144150069e-06, 2.447966687239273e-06, 2.3677339844768586e-06, 2.2901309280352534e-06, 2.2150713306175765e-06, 2.1424718297365362e-06, 2.0722517951306097e-06, 2.00433323921468e-06, 1.9386407304656755e-06, 1.8751013096470136e-06, 1.8136444087788125e-06, 1.7542017727638716e-06, 1.6967073835823792e-06, 1.6410973869711579e-06, 1.587310021506016e-06, 1.5352855500084407e-06, 1.4849661932004546e-06, 1.43629606553395e-06, 1.3892211131232327e-06, 1.343689053711841e-06, 1.2996493186069677e-06, 1.2570529965169952e-06, 1.2158527792297687e-06, 1.1760029090712773e-06, 1.1374591280863902e-06, 1.1001786288852058e-06, 1.0641200071004238e-06, 1.0292432154029389e-06, 9.955095190245847e-07, 9.628814527386294e-07, 9.313227792502477e-07, 9.007984489507533e-07, 8.712745609908984e-07, 8.427183256300032e-07, 8.150980278191031e-07, 7.883829919776665e-07, 7.625435479247645e-07, 7.375509979268551e-07, 7.13377584825583e-07, 6.899964612102005e-07, 6.673816596003678e-07, 6.455080636062205e-07, 6.243513800336727e-07, 6.038881119039751e-07, 5.840955323575642e-07, 5.64951659413219e-07, 5.464352315544918e-07, 5.28525684116301e-07, 5.11203126445458e-07, 4.944483198097636e-07, 4.782426560311396e-07, 4.625681368190636e-07, 4.474073537813572e-07, 4.3274346909012354e-07, 4.1856019678136516e-07, 4.048417846675101e-07, 3.915729968427603e-07, 3.7873909676183077e-07, 3.6632583087328767e-07, 3.543194127893074e-07, 3.427065079742759e-07, 3.314742189352227e-07, 3.2061007089764216e-07, 3.1010199795079294e-07, 2.9993832964708907e-07, 2.901077780406988e-07, 2.8059942515095677e-07, 2.7140271083666573e-07, 2.6250742106782124e-07, 2.539036765817331e-07, 2.455819219109453e-07, 2.3753291477076871e-07, 2.2974771579463978e-07, 2.2221767860590528e-07, 2.1493444021500694e-07, 2.078899117314006e-07, 2.0107626937989435e-07, 1.9448594581142852e-07, 1.8811162169864695e-07, 1.8194621760692534e-07, 1.7598288613182872e-07, 1.7021500429426566e-07, 1.6463616618489314e-07, 1.5924017584960275e-07, 1.5402104040818694e-07, 1.4897296339854257e-07, 1.4409033833901985e-07, 1.3936774250176683e-07, 1.3479993089015402e-07, 1.3038183041359041e-07, 1.2610853425326133e-07, 1.2197529641253058e-07, 1.179775264459545e-07, 1.1411078436105377e-07, 1.1037077568718104e-07, 1.0675334670600752e-07, 1.0325447983833153e-07, 9.987028918208557e-08, 9.659701619658623e-08, 9.343102552823394e-08, 9.036880097302633e-08, 8.740694157140127e-08, 8.454215783107235e-08, 8.17712680736618e-08, 7.909119490107352e-08, 7.649896177768146e-08, 7.399168972453768e-08, 7.156659412192854e-08, 6.92209816167279e-08, 6.695224713111248e-08, 6.475787096931743e-08, 6.263541601921861e-08, 6.058252504563368e-08, 5.859691807233595e-08, 5.6676389849873365e-08, 5.481880740638028e-08, 5.3022107678662035e-08, 5.128429522092127e-08, 4.96034399885813e-08, 4.797767519474524e-08, 4.640519523691018e-08, 4.48842536916338e-08, 4.3413161374926284e-08, 4.1990284466213385e-08, 4.0614042693787e-08, 3.928290757972811e-08, 3.799540074235274e-08, 3.6750092254295673e-08, 3.5545599054408434e-08, 3.4380583411707595e-08, 3.325375143966765e-08, 3.2163851659208233e-08, 3.110967360877981e-08, 3.009004650000413e-08, 2.910383791737644e-08, 2.8149952560585187e-08, 2.722733102805257e-08, 2.6334948640344825e-08, 2.5471814302145515e-08, 2.4636969401527923e-08, 2.3829486745304066e-08, 2.304846952926791e-08, 2.2293050342189097e-08, 2.1562390202451026e-08, 2.0855677626263363e-08, 2.0172127726414107e-08, 1.951098134056027e-08, 1.887150418808905e-08, 1.825298605461307e-08, 1.7654740003194024e-08, 1.707610161141859e-08, 1.6516428233479435e-08, 1.597509828644165e-08, 1.5451510559901994e-08, 1.4945083548274222e-08, 1.4455254804958923e-08, 1.3981480317680589e-08, 1.3523233904298183e-08, 1.308000662841814e-08, 1.2651306234160813e-08, 1.2236656599452574e-08, 1.1835597207236404e-08, 1.1447682634013684e-08, 1.107248205514915e-08, 1.0709578766389603e-08, 1.0358569721064955e-08, 1.0019065082457628e-08, 9.690687790843147e-09, 9.373073144721095e-09, 9.065868395771311e-09, 8.768732357085509e-09, 8.48133502423919e-09, 8.203357208783027e-09, 7.934490183746643e-09, 7.674435340761104e-09, 7.4229038584192995e-09, 7.179616381505885e-09, 6.944302710740554e-09, 6.71670150269004e-09, 6.49655997951558e-09, 6.283633648233474e-09, 6.077686029176948e-09, 5.8784883933577446e-09, 5.6858195084357525e-09, 5.499465393014544e-09, 5.319219078989932e-09, 5.144880381687616e-09, 4.976255677534617e-09, 4.813157689017592e-09, 4.655405276689185e-09, 4.502823237991422e-09, 4.355242112672724e-09, 4.212497994582418e-09, 4.074432349633728e-09, 3.940891839733082e-09, 3.811728152480167e-09, 3.6867978364496143e-09, 3.565962141871365e-09, 3.449086866532775e-09, 3.336042206731315e-09, 3.226702613112335e-09, 3.1209466512317793e-09, 3.01865686668899e-09, 2.9197196546798196e-09, 2.824025133825171e-09, 2.7314670241348346e-09, 2.6419425289710957e-09, 2.5553522208810084e-09, 2.4715999311705473e-09, 2.3905926430979923e-09, 2.312240388567926e-09, 2.2364561482111102e-09, 2.16315575473927e-09, 2.0922577994674476e-09, 2.0236835419001073e-09, 1.9573568222805806e-09, 1.893203977006722e-09, 1.8311537568188385e-09, 1.7711372476690293e-09, 1.7130877941840522e-09, 1.6569409256367131e-09, 1.6026342843435607e-09, 1.5501075564093651e-09, 1.4993024047414625e-09, 1.4501624042595701e-09, 1.4026329792291169e-09, 1.3566613426484883e-09, 1.3121964376228692e-09, 1.269188880659573e-09, 1.2275909068218811e-09, 1.1873563166804771e-09, 1.1484404250035622e-09, 1.1108000111286633e-09, 1.0743932709610178e-09, 1.0391797705452224e-09, 1.0051204011585838e-09, 9.721773358762935e-10, 9.40313987560192e-10, 9.09494968224459e-10, 8.796860497331055e-10, 8.508541257856127e-10, 8.229671751485029e-10, 7.959942260920032e-10, 7.699053219923065e-10, 7.446714880612275e-10, 7.202646991663012e-10, 6.966578487055862e-10, 6.738247185026036e-10, 6.51739949687977e-10, 6.303790145354329e-10, 6.097181892208843e-10, 5.897345274743407e-10, 5.70405835095384e-10, 5.517106453039058e-10, 5.336281948987295e-10, 5.161384011976404e-10, 4.992218397332114e-10, 4.828597226796531e-10, 4.670338779867299e-10, 4.517267291975665e-10, 4.3692127592793087e-10, 4.2260107498531325e-10, 4.087502221068324e-10, 3.95353334295686e-10, 3.823955327365283e-10, 3.6986242627080047e-10, 3.577400954136609e-10, 3.460150768947645e-10, 3.3467434870572196e-10, 3.237053156376318e-10, 3.1309579529262354e-10, 3.028340045538759e-10, 2.929085464990834e-10, 2.833083977428374e-10, 2.7402289619386344e-10, 2.650417292135184e-10, 2.5635492216239534e-10, 2.479528273223168e-10, 2.3982611318141184e-10, 2.3196575407037768e-10, 2.2436302013841518e-10, 2.1700946765770551e-10, 2.0989692964566e-10, 2.0301750679452808e-10, 1.9636355869828928e-10, 1.899276953670863e-10, 1.837027690197743e-10, 1.776818661454717e-10, 1.718582998252952e-10, 1.6622560230575213e-10, 1.6077751781554145e-10, 1.5550799561778607e-10, 1.5041118328997963e-10, 1.4548142022418496e-10, 1.4071323134026492e-10, 1.361013210051637e-10, 1.316405671514852e-10, 1.2732601558883627e-10, 1.2315287450161727e-10, 1.191165091271487e-10, 1.1521243660822361e-10, 1.114363210143689e-10, 1.0778396852628582e-10, 1.0425132277812184e-10, 1.0083446035240045e-10, 9.752958642260591e-11, 9.433303053858326e-11, 9.12412425500729e-11, 8.825078866385227e-11, 8.535834763010572e-11, 8.2560707053787e-11, 7.985475982687779e-11, 7.723750067757973e-11, 7.470602283260761e-11, 7.225751478887667e-11, 6.988925719099866e-11, 6.759861981111865e-11, 6.538305862773854e-11, 6.324011300028275e-11, 6.116740293626822e-11, 5.916262644804359e-11, 5.722355699616193e-11, 5.534804101654747e-11, 5.353399552871012e-11, 5.177940582235127e-11, 5.008232321979171e-11, 4.844086291173649e-11, 4.6853201863973094e-11, 4.5317576792678116e-11, 4.3832282206083756e-11, 4.239566851032915e-11, 4.100614017739289e-11, 3.966215397307202e-11, 3.836221724303941e-11, 3.7104886255076053e-11, 3.5888764595637097e-11, 3.47125016189708e-11, 3.357479094706798e-11, 3.247436901877603e-11, 3.1410013686466016e-11, 3.038054285869439e-11, 2.938481318735178e-11, 2.8421718797840827e-11, 2.7490190060872687e-11, 2.6589192404518272e-11, 2.5717725165194752e-11, 2.4874820476311277e-11, 2.4059542193339603e-11, 2.3270984854115772e-11, 2.2508272673218182e-11, 2.1770558569305145e-11, 2.1057023224331696e-11, 2.0366874173600794e-11, 1.9699344925638328e-11, 1.9053694110914423e-11, 1.8429204658465624e-11, 1.7825182999503465e-11, 1.7240958297124986e-11, 1.667588170126966e-11, 1.6129325628095298e-11, 1.560068306297257e-11, 1.5089366886324058e-11, 1.4594809221559104e-11, 1.411646080438024e-11, 1.3653790372760758e-11, 1.3206284076915912e-11, 1.2773444908612462e-11, 1.2354792149182733e-11, 1.1949860835630141e-11, 1.1558201244233253e-11, 1.117937839107484e-11, 1.0812971548941211e-11, 1.0458573780055297e-11, 1.0115791484124512e-11, 9.784243961201463e-12, 9.463562988872001e-12, 9.153392413301037e-12, 8.853387753681936e-12, 8.563215819650167e-12, 8.282554341236332e-12, 8.01109161094755e-12, 7.748526137579734e-12, 7.494566311376238e-12, 7.2489300801610225e-12, 7.011344636086622e-12, 6.7815461126490356e-12, 6.559279291633022e-12, 6.344297319662334e-12, 6.1363614340400904e-12, 5.935240697574792e-12, 5.740711742097484e-12, 5.552558520385204e-12, 5.3705720662151955e-12, 5.1945502622834124e-12, 5.024297615729538e-12, 4.859625041019239e-12, 4.700349649942493e-12, 4.5462945484947855e-12, 4.397288640415567e-12, 4.253166437165781e-12, 4.113767874133439e-12, 3.978938132863085e-12, 3.848527469111746e-12, 3.722391046540385e-12, 3.600388775856154e-12, 3.4823851592268056e-12, 3.368249139794456e-12, 3.25785395612157e-12, 3.1510770014075175e-12, 3.047799687319339e-12, 2.947907312285491e-12, 2.8512889341062987e-12, 2.757837246739627e-12, 2.6674484611249365e-12, 2.5800221899133554e-12, 2.4954613359757553e-12, 2.413671984564998e-12, 2.3345632990125947e-12, 2.2580474198439294e-12, 2.1840393672000064e-12, 2.1124569464573497e-12, 2.0432206569412322e-12, 1.9762536036308506e-12, 1.911481411758386e-12, 1.8488321442070994e-12, 1.7882362216167256e-12, 1.7296263451074323e-12, 1.6729374215365205e-12, 1.6181064912048532e-12, 1.565072657932725e-12, 1.5137770214275113e-12]
F=numpy.array(RK4)

analytic=numpy.array([1000,
  967.2247567687975,
  935.5237301064595,
  904.8617123036585,
  875.2046495923036,
  846.5196043248365,
  818.7747183931086,
  791.9391778462151,
  765.983178667987,
  740.877893676134,
  716.5954405062777,
  693.1088506453139,
  670.3920394797144,
  648.4197773255048,
  627.1676614077393,
  606.6120887583561,
  586.7302300023132,
  567.5000040028879],
 967.2247567687975,
 935.5237301064595,
 904.8617123036585,
 875.2046495923036,
 846.5196043248365,
 818.7747183931086,
 791.9391778462151,
 765.983178667987,
 740.877893676134,
 716.5954405062777,
 693.1088506453139,
 670.3920394797144,
 648.4197773255048,
 627.1676614077393,
 606.6120887583561,
 586.7302300023132,
 567.5000040028879,
 548.9000533379849,
 530.9097205802124,
 513.5090253543862,
 496.67864214697846,
 480.39987884286785,
 464.6546559655526,
 449.42548659777094,
 434.69545696022743,
 420.44820762685725,
 406.66791535576385,
 393.3392755156526,
 380.4474850882421,
 367.9782262277757,
 355.9176503593739,
 344.25236279856733,
 332.96940787492815,
 322.0562545432779,
 311.5007824664919,
 301.29126855444264,
 291.4163739441333,
 281.8651314065592,
 272.62693316631436,
 263.6915191204116,
 255.04896544323492,
 246.68967356496628,
 238.6043595112486,
 230.7840435922421,
 223.22004042962593,
 215.9039493104661,
 208.8276448572383,
 201.98326800364316,
 195.3632172661906,
 188.96014030186092,
 182.76692574246528,
 176.7766952966369,
 170.98279611068142,
 165.37879337980274,
 159.95846320149693,
 154.7157856631785,
 149.6449381563612,
 144.7402889099682,
 139.99639073558947,
 135.40797497774008,
 130.96994566240008,
 126.67737383733754,
 122.52549209792883,
 118.50968929239643,
 114.6255054005839,
 110.86862658058024,
 107.23488037769235,
 103.72023109044459,
 100.32077528845873,
 97.03273747723668,
 93.85246590503083,
 90.77642850714528,
 87.80120898316373,
 84.92350300274691,
 82.14011453578613,
 79.4479523028369,
 76.84402634189044,
 74.32544468767004,
 71.88941015976437,
 69.53321725603041,
 67.254249147816,
 65.04997477366442,
 62.91794602827398,
 60.85579504358963,
 58.86123155900777,
 56.932040377773156,
 55.066078906742995,
 53.2612747767859,
 51.51562354117283,
 49.82718644940383,
 48.19408829399814,
 46.6145153278563,
 45.08671324988119,
 43.608985256620855,
 42.17969015776919,
 40.797240553431536,
 39.460101071130936,
 38.166786660596784,
 36.91586094444231,
 35.705934622898965,
 34.53566393083604,
 33.40374914535182,
 32.30893314227883,
 31.25,
 30.225773649024923,
 29.235116565826857,
 28.276928509489323,
 27.350145299759482,
 26.45373763515114,
 25.58670994978464,
 24.74809930769422,
 23.93697433337459,
 23.152434177379185,
 22.393607515821174,
 21.65965158266605,
 20.949751233741072,
 20.26311804142202,
 19.598989418991845,
 18.956627773698624,
 18.33531968757228,
 17.73437512509025,
 17.15312666681203,
 16.59092876813164,
 16.04715704232457,
 15.521207567093077,
 15.01249621383962,
 14.520457998923519,
 14.044546456180342,
 13.584233030007107,
 13.13900648833929,
 12.70837235486762,
 12.291852359864144,
 11.888983909007566,
 11.49931956961799,
 11.122426573730433,
 10.757886337455226,
 10.405293996091503,
 10.064257954477432,
 9.734399452077868,
 9.41535214232633,
 9.106761685754162,
 8.808285356454974,
 8.519591661447322,
 8.240359972512861,
 7.970280170101088,
 7.709052298905193,
 7.456386234726515,
 7.2120013622575625,
 6.975626263425807,
 6.746998415952062,
 6.525863901788695,
 6.311977125113846,
 6.105100539568453,
 5.905004384433151,
 5.711466429452037,
 5.524271728019898,
 5.343212378458791,
 5.168087293118832,
 4.998701975046775,
 4.834868301974325,
 4.676404317386285,
 4.523134028436508,
 4.374887210487172,
 4.2314992180543785,
 4.0928108019500025,
 3.958667932416798,
 3.828921628060276,
 3.7034277903873885,
 3.582047043768247,
 3.464644580643133,
 3.351090011802886,
 3.2412572215763933,
 3.1350242277643354,
 3.0322730461636462,
 2.9328895595322133,
 2.8367633908482905,
 2.7437877807238666,
 2.653859468835841,
 2.5668785792433164,
 2.482748509463653,
 2.4013758231840763,
 2.322670146489689,
 2.2465440674926365,
 2.1729130392509504,
 2.1016952858692486,
 2.032811711677012,
 1.966185813383561,
 1.9017435951121753,
 1.8394134862189921,
 1.7791262618054104,
 1.720814965835718,
 1.6644148367745588,
 1.6098632356616505,
 1.557099576543869,
 1.5060652591874413,
 1.456703603995509,
 1.4089597890587866,
 1.3627807892694013,
 1.3181153174302864,
 1.274913767294735,
 1.2331281584728413,
 1.192712083143649,
 1.1536206545138217,
 1.1158104569655922,
 1.0792394978386257,
 1.0438671607922438,
 1.009654160696213,
 0.9765625,
 0.9445554265320288,
 0.9135973926820893,
 0.8836540159215414,
 0.8546920406174838,
 0.8266793010984731,
 0.79958468593077,
 0.7733781033654443,
 0.7480304479179559,
 0.7235135680430995,
 0.6998002348694117,
 0.6768641119583141,
 0.6546797260544085,
 0.6332224387944382,
 0.6124684193434952,
 0.592394617928082,
 0.5729787402366338,
 0.5541992226590701,
 0.5360352083378758,
 0.5184665240041135,
 0.5014736575726426,
 0.4850377364716585,
 0.46914050668248797,
 0.4537643124663598,
 0.43889207675563546,
 0.4245072821877219,
 0.41059395276060257,
 0.3971366360896129,
 0.3841203862457543,
 0.37153074715648626,
 0.359353736550562,
 0.34757583042907586,
 0.33618394804547574,
 0.3251654373778593,
 0.31450806107741985,
 0.3041999828774333,
 0.294229754447698,
 0.28458630267981744,
 0.27525891738921804,
 0.2662372394202287,
 0.257511249141027,
 0.24907125531565888,
 0.2409078843407874,
 0.23301206983520348,
 0.22537504257054894,
 0.21798832073205637,
 0.21084370049850204,
 0.20393324693089662,
 0.19724928515980777,
 0.19078439186151408,
 0.18453138701353605,
 0.17848332592037608,
 0.1726334915006219,
 0.16697538682683713,
 0.1615027279099636,
 0.15620943672021165,
 0.1510896344366977,
 0.14613763491832135,
 0.1413479383886408,
 0.13671522532772395,
 0.13223435056419927,
 0.1279003375609374,
 0.12370837288802489,
 0.11965380087688346,
 0.11573211844960583,
 0.11193897011775757,
 0.10827014314509785,
 0.10472156286884006,
 0.10128928817426225,
 0.09796950711763552,
 0.0947585326926139,
 0.09165279873538171,
 0.08864885596400905,
 0.08574336814762087,
 0.08293310840111999,
 0.08021495560135367,
 0.07758589092073911,
 0.07504299447450241,
 0.07258344207780275,
 0.07020450210914492,
 0.06790353247659216,
 0.06567797768341406,
 0.0635253659899066,
 0.06144330666823631,
 0.05942948734725545,
 0.057481671444343525,
 0.05559769568141905,
 0.053775467682366206,
 0.05201296364920494,
 0.0503082261144266,
 0.04865936176699589,
 0.04706453934960756,
 0.045521987624859636,
 0.0440299934080871,
 0.04258689966466877,
 0.04119110366969647,
 0.03984105522796045,
 0.038535254952276304,
 0.03727225259823901,
 0.03605064545355694,
 0.03486907678017474,
 0.03372623430745707,
 0.0326208487747576,
 0.03155169252175668,
 0.03051757812499996,
 0.029517357079125887,
 0.028549918521315252,
 0.027614187997548157,
 0.02670912626929634,
 0.025833728159327274,
 0.024987021435336532,
 0.024168065730170125,
 0.02337595149743609,
 0.02260979900134685,
 0.021868757339669087,
 0.021152003498697308,
 0.020458741439200272,
 0.019788201212326183,
 0.019139638104484234,
 0.018512331810252556,
 0.017905585632394815,
 0.017318725708095936,
 0.016751100260558625,
 0.01620207887512854,
 0.01567105179914509,
 0.01515742926473932,
 0.014660640833827754,
 0.014180134764573736,
 0.013715377398613615,
 0.013265852568366304,
 0.012831061023768837,
 0.012410519877800398,
 0.012003762070179828,
 0.01161033584864019,
 0.011229804267205068,
 0.010861744700908615,
 0.010505748376421119,
 0.010161419918058098,
 0.009828376908669367,
 0.009506249464919785,
 0.009194679826490557,
 0.008893321958744291,
 0.00860184116841306,
 0.008319913731882143,
 0.008047226535657091,
 0.0077834767286143365,
 0.007528371385649602,
 0.007281627182350105,
 0.007042970080329651,
 0.006812135022876758,
 0.006588865640578185,
 0.006372913966590516,
 0.00616404016124399,
 0.005962012245672312,
 0.005766605844172999,
 0.00557760393501175,
 0.005394796609394432,
 0.0052179808383386594,
 0.0050469602471863596,
 0.004881544897506612,
 0.004721551076146802,
 0.00456680109119754,
 0.004417123074645023,
 0.00427235079149138,
 0.0041323234551312255,
 0.003996885548779299,
 0.003865886652750776,
 0.0037391812774026133,
 0.003616628701550181,
 0.0034980928161799287,
 0.003383441973284306,
 0.0032725488396512558,
 0.003165290255445694,
 0.003061547097426109,
 0.002961204146644183,
 0.002864149960480677,
 0.0027702767488752816,
 0.0026794802546131515,
 0.0025916596375349984,
 0.002506717362542301,
 0.0024245590912730965,
 0.0023450935773281995,
 0.002268232564931335,
 0.002193890690910778,
 0.002121985389893504,
 0.002052436802606688,
 0.0019851676871845804,
 0.0019201033333823838,
 0.001857171479601732,
 0.0017963022326357345,
 0.0017374279900443444,
 0.0016804833650739433,
 0.0016254051140376536,
 0.0015721320660758305,
 0.0015206050552186209,
 0.0014707668546752356,
 0.001422562113276863,
 0.0013759372940027213,
 0.0013308406145208984,
 0.001287221989678014,
 0.0012450329758737636,
 0.001204226717258634,
 0.0011647578936949687,
 0.001126582670423654,
 0.00108965864938046,
 0.001053944822108033,
 0.0010194015242111745,
 0.0009859903913048956,
 0.00095367431640625,
 0.0009224174087226835,
 0.0008921849537911028,
 0.0008629433749233795,
 0.0008346601959155117,
 0.0008073040049789769,
 0.0007808444198542676,
 0.0007552520540678162,
 0.0007304984842948788,
 0.0007065562187920887,
 0.0006833986668646598,
 0.0006610001093342906,
 0.0006393356699750083,
 0.0006183812878851929,
 0.000598113690765132,
 0.000578510369070392,
 0.0005595495510123378,
 0.0005412101783779977,
 0.0005234718831424568,
 0.0005063149648477667,
 0.0004897203687232838,
 0.00047366966452310356,
 0.00045814502605711716,
 0.0004431292113929291,
 0.00042860554370667526,
 0.00041455789276144677,
 0.00040097065699277594,
 0.0003878287461812623,
 0.0003751175646931194,
 0.0003628229952700058,
 0.0003509313833501582,
 0.0003394295219033941,
 0.0003283046367631599,
 0.00031754437243931545,
 0.00030713677839591755,
 0.00029707029577874316,
 0.0002873337445778298,
 0.00027791631121075895,
 0.000268807536512908,
 0.00025999730412131687,
 0.000251475829239284,
 0.0002432336477691979,
 0.00023526160580155,
 0.00022755084944844068,
 0.0002200928150103015,
 0.0002128792194648986,
 0.0002059020512680682,
 0.0001991535614559539,
 0.0001926262550388746,
 0.00018631288267726,
 0.00018020643263040614,
 0.00017430012296911743,
 0.00016858739404357594,
 0.0001630619011980833,
 0.00015771750772457366,
 0.00015254827804708182,
 0.00014754847112958747,
 0.0001427125340999233,
 0.00013803509608265692,
 0.00013351096223410553,
 0.00012913510797285074,
 0.00012490267339935305,
 0.0001208089578984617,
 0.00011684941491883161,
 0.0001130196469234431,
 0.00010931540050562272,
 0.00010573256166513452,
 0.00010226715123910153,
 9.891532048267806e-05,
 9.567334679456586e-05,
 9.253762958263068e-05,
 8.950468626502095e-05,
 8.657114840235266e-05,
 8.373375795666094e-05,
 8.098936367296867e-05,
 7.833491757944672e-05,
 7.576747160228436e-05,
 7.32841742915062e-05,
 7.088226765410418e-05,
 6.855908409096164e-05,
 6.63120434341721e-05,
 6.413865008145898e-05,
 6.203649022451811e-05,
 6.0003229168199365e-05,
 5.8036608737554204e-05,
 5.6134444769866677e-05,
 5.429462468888574e-05,
 5.251510515856061e-05,
 5.079390981367674e-05,
 4.9129127064869684e-05,
 4.751890797558188e-05,
 4.596146420860101e-05,
 4.445506603990203e-05,
 4.299804043758503e-05,
 4.1588769203778054e-05,
 4.0225687177437994e-05,
 3.890728049605516e-05,
 3.76320849143323e-05,
 3.639868417796776e-05,
 3.5205708450739234e-05,
 3.4051832793139425e-05,
 3.2935775690876015e-05,
 3.185629763159919e-05,
 3.0812199728278034e-05,
 2.98023223876953e-05,
 2.8825544022583847e-05,
 2.7880779805971902e-05,
 2.6966980466355646e-05,
 2.6083131122359726e-05,
 2.5228250155593017e-05,
 2.4401388120445812e-05,
 2.3601626689619286e-05,
 2.2828077634214953e-05,
 2.2079881837252762e-05,
 2.1356208339520573e-05,
 2.0656253416696607e-05,
 1.9979239686719e-05,
 1.932441524641227e-05,
 1.8691052836410334e-05,
 1.8078449033449775e-05,
 1.7485923469135545e-05,
 1.691281807431242e-05,
 1.6358496348201738e-05,
 1.582234265149273e-05,
 1.5303761522602612e-05,
 1.480217701634698e-05,
 1.4317032064284877e-05,
 1.3847787856029053e-05,
 1.3393923240833597e-05,
 1.2954934148795206e-05,
 1.2530333031024221e-05,
 1.2119648318164463e-05,
 1.1722423896659977e-05,
 1.1338218602187676e-05,
 1.096660572969242e-05,
 1.060717255948108e-05,
 1.0259519898848742e-05,
 9.923261638728603e-06,
 9.598024324872437e-06,
 9.283446743085736e-06,
 8.979179518057178e-06,
 8.684884725336214e-06,
 8.400235516028386e-06,
 8.124915753791162e-06,
 7.858619663727621e-06,
 7.601051492787431e-06,
 7.351925181298447e-06,
 7.110964045263781e-06,
 6.877900469071919e-06,
 6.652475608278078e-06,
 6.43443910212714e-06,
 6.223548795498557e-06,
 6.019570469964828e-06,
 5.8222775836643625e-06,
 5.6314510197001995e-06,
 5.446878842784917e-06,
 5.268356063861746e-06,
 5.095684412440092e-06,
 4.9286721163929335e-06,
 4.767133688971305e-06,
 4.610889722799607e-06,
 4.459766690622593e-06,
 4.3135967525830345e-06,
 4.172217569815796e-06,
 4.035472124151584e-06,
 3.903208543729774e-06,
 3.775279934326933e-06,
 3.651544216213486e-06,
 3.5318639663575952e-06,
 3.4161062658007023e-06,
 3.3041425520354585e-06,
 3.1958484762219273e-06,
 3.0911037650836826e-06,
 2.9897920873301763e-06,
 2.8918009244572127e-06,
 2.7970214457819086e-06,
 2.7053483875735146e-06,
 2.616679936145649e-06,
 2.5309176147802744e-06,
 2.4479661743577135e-06,
 2.367733487571381e-06,
 2.2901304466095635e-06,
 2.215070864190759e-06,
 2.1424713778425543e-06,
 2.0722513573178733e-06,
 2.004332815045596e-06,
 1.9386403195161935e-06,
 1.8751009115062327e-06,
 1.8136440230485646e-06,
 1.754201399058336e-06,
 1.6967070215276818e-06,
 1.6410970362050215e-06,
 1.5873096816773946e-06,
 1.5352852207771795e-06,
 1.4849658742369357e-06,
 1.4362957565187835e-06,
 1.3892208137469352e-06,
 1.3436887636745338e-06,
 1.299649037618066e-06,
 1.2570527242949346e-06,
 1.215852515501721e-06,
 1.176002653572886e-06,
 1.137458880561494e-06,
 1.1001783890855987e-06,
 1.0641197747856047e-06,
 1.0292429903398768e-06,
 9.955093009874762e-07,
 9.628812415086862e-07,
 9.313225746154761e-07,
 9.007982507057465e-07,
 8.712743689366231e-07,
 8.42718139573612e-07,
 8.150978475737397e-07,
 7.883828173622828e-07,
 7.625433787639327e-07,
 7.375508340506011e-07,
 7.133774260692157e-07,
 6.899963074141498e-07,
 6.673815106100188e-07,
 6.455079192717675e-07,
 6.243512402099673e-07,
 6.038879764503843e-07,
 5.840954011378237e-07,
 5.649515322953042e-07,
 5.464351084104846e-07,
 5.285255648222638e-07,
 5.11203010881305e-07,
 4.944482078591467e-07,
 4.782425475813306e-07,
 4.6256803176084374e-07,
 4.4740725200890305e-07,
 4.3274337050090695e-07,
 4.185601012760504e-07,
 4.0484169214985073e-07,
 3.9157290721950744e-07,
 3.7873900994263864e-07,
 3.663257467706248e-07,
 3.5431933131836535e-07,
 3.4270642905288857e-07,
 3.31474142483783e-07,
 3.2060999683902356e-07,
 3.1010192621026926e-07,
 2.9993826015226297e-07,
 2.901077107214286e-07,
 2.805993599392872e-07,
 2.7140264766675705e-07,
 2.625073598758865e-07,
 2.5390361730597324e-07,
 2.455818644914885e-07,
 2.3753285914960754e-07,
 2.2974766191557598e-07,
 2.2221762641449267e-07,
 2.1493438965849775e-07,
 2.0788986275869024e-07,
 2.0107622194147267e-07,
 1.944858998593295e-07,
 1.8811157718640115e-07,
 1.8194617448951157e-07,
 1.7598284436563084e-07,
 1.702149638370283e-07,
 1.6463612699567976e-07,
 1.5924013788875306e-07,
 1.5402100363727883e-07,
 1.4897292778035295e-07,
 1.440903038374879e-07,
 1.3936770908195625e-07,
 1.347998985182195e-07,
 1.3038179905674334e-07,
 1.2610850387973716e-07,
 1.219752669915556e-07,
 1.179774979477164e-07,
 1.141107567566712e-07,
 1.1037074894867501e-07,
 1.0675332080627209e-07,
 1.0325445475110784e-07,
 9.987026488193499e-08,
 9.659699265886521e-08,
 9.343100272906813e-08,
 9.03687788892877e-08,
 8.740692018068476e-08,
 8.454213711167245e-08,
 8.177124800455162e-08,
 7.909117546188339e-08,
 7.649894294867865e-08,
 7.399167148660575e-08,
 7.156657645654895e-08,
 6.922096450596106e-08,
 6.695223055757992e-08,
 6.475785491618362e-08,
 6.263540047017473e-08,
 6.058250998488093e-08,
 5.8596903484569845e-08,
 5.6676375720267725e-08,
 5.4818793720572873e-08,
 5.302209442273994e-08,
 5.128428238140698e-08,
 4.9603427552418654e-08,
 4.7977663149286754e-08,
 4.640518356990414e-08,
 4.4884242391212044e-08,
 4.3413150429591786e-08,
 4.199027386482909e-08,
 4.061403242556447e-08,
 3.928289763421676e-08,
 3.7995391109428834e-08,
 3.6750082924152605e-08,
 3.554559001754661e-08,
 3.4380574658925e-08,
 3.325374296205019e-08,
 3.2163843448121076e-08,
 3.110966565585856e-08,
 3.009003879714649e-08,
 2.9103830456733664e-08,
 2.8149945334554515e-08,
 2.722732402926941e-08,
 2.6334941861675412e-08,
 2.54718077366794e-08,
 2.4636963042571283e-08,
 2.382948058637284e-08,
 2.304846356408131e-08,
 2.229304456466302e-08,
 2.156238460669213e-08,
 2.085567220656304e-08,
 2.0172122477242764e-08,
 1.9510976256561506e-08,
 1.8871499264074467e-08,
 1.8252981285557017e-08,
 1.7654735384228283e-08,
 1.7076097137827667e-08,
 1.651642390069571e-08,
 1.5975094090040803e-08,
 1.5451506495598355e-08,
 1.49450796119166e-08,
 1.4455250992526334e-08,
 1.398147662527824e-08,
 1.3523230328153359e-08,
 1.3080003164876546e-08,
 1.2651302879682809e-08,
 1.2236653350609624e-08,
 1.1835594060707472e-08,
 1.1447679586581999e-08,
 1.1072479103698894e-08,
 1.0709575907902783e-08,
 1.0358566952618234e-08,
 1.0019062401219465e-08,
 9.690685194070893e-09,
 9.373070629758231e-09,
 9.065865960044657e-09,
 8.768729998102706e-09,
 8.48133273958614e-09,
 8.203354996121464e-09,
 7.934488040811675e-09,
 7.674433265358998e-09,
 7.422901848425219e-09,
 7.1796144348617585e-09,
 6.944300825452905e-09,
 6.71669967682804e-09,
 6.496558211209056e-09,
 6.283631935671029e-09,
 6.077684370604054e-09,
 5.878486787075023e-09,
 5.685817952797224e-09,
 5.499463886425971e-09,
 5.319217619907141e-09,
 5.144878968614982e-09,
 4.976254309023523e-09,
 4.81315636366497e-09,
 4.655403993136036e-09,
 4.502821994921487e-09,
 4.3552409088111385e-09,
 4.2124968286943655e-09,
 4.074431220523235e-09,
 3.940890746241777e-09,
 3.811727093486118e-09,
 3.6867968108661422e-09,
 3.5659611486459797e-09,
 3.4490859046460862e-09,
 3.3360412751960075e-09,
 3.2267017109721246e-09,
 3.120945777560473e-09,
 3.0186560205895313e-09,
 2.919718835283383e-09,
 2.8240243402902443e-09,
 2.7314662556463927e-09,
 2.641941784739758e-09,
 2.555351500142242e-09,
 2.4715992331838592e-09,
 2.3905919671462025e-09,
 2.312239733956425e-09,
 2.236455514267158e-09,
 2.163155140811286e-09,
 2.092257204924368e-09,
 2.0236829661307336e-09,
 1.957356264692963e-09,
 1.893203437027531e-09,
 1.8311532338928037e-09,
 1.7711367412583623e-09,
 1.7130873037679048e-09,
 1.6569404507106253e-09,
 1.6026338244189648e-09,
 1.5501071110130796e-09,
 1.4993019734152131e-09,
 1.4501619865595063e-09,
 1.4026325747253733e-09,
 1.3566609509247402e-09,
 1.3121960582759109e-09,
 1.2691885132988914e-09,
 1.2275905510692708e-09,
 1.1873559721696484e-09,
 1.1484400913797704e-09,
 1.1107996880483329e-09,
 1.074392958091404e-09,
 1.039179467564066e-09,
 1.005120107753785e-09,
 9.721770517455812e-10,
 9.403137124108256e-10,
 9.094947017729282e-10,
 8.796857917048297e-10,
 8.508538759146703e-10,
 8.229669331773548e-10,
 7.959939917712323e-10,
 7.699050950803536e-10,
 7.446712683241523e-10,
 7.202644863775395e-10,
 6.966576426457203e-10,
 6.7382451895913e-10,
 6.517397564550959e-10,
 6.303788274138349e-10,
 6.097180080175479e-10,
 5.897343520023279e-10,
 5.704056651736554e-10,
 5.517104807571326e-10,
 5.336280355571153e-10,
 5.161382468967416e-10,
 4.992216903137739e-10,
 4.828595779874475e-10,
 4.670337378723944e-10,
 4.5172659351644855e-10,
 4.3692114453994397e-10,
 4.226009477547916e-10,
 4.0875009890239263e-10,
 3.9535321499008824e-10,
 3.8239541720654987e-10,
 3.6986231439710773e-10,
 3.57739987080688e-10,
 3.4601497199059085e-10,
 3.3467424712196124e-10,
 3.2370521726931903e-10,
 3.1309570003810874e-10,
 3.028339123147158e-10,
 2.9290845717994405e-10,
 2.8330831125139486e-10,
 2.740228124407099e-10,
 2.650416481120672e-10,
 2.563548436287952e-10,
 2.4795275127536434e-10,
 2.39826039542469e-10,
 2.319656827632884e-10,
 2.2436295108942946e-10,
 2.1700940079540281e-10,
 2.0989686490087653e-10,
 2.0301744410028325e-10,
 1.9636349798971923e-10,
 1.8992763658137696e-10,
 1.8370271209609472e-10,
 1.7768181102491348e-10,
 1.7185824645081123e-10,
 1.6622555062209838e-10,
 1.607774677692184e-10,
 1.555079471569853e-10,
 1.5041113636452998e-10,
 1.4548137478550134e-10,
 1.4071318734129667e-10,
 1.3610127840034777e-10,
 1.3164052589669866e-10,
 1.2732597564135125e-10,
 1.2315283582005572e-10,
 1.1911647167144092e-10,
 1.152124003395667e-10,
 1.1143628589518701e-10,
 1.0778393452019034e-10,
 1.04251289849875e-10,
 1.0083442846787866e-10,
 9.75295555487649e-11,
 9.433300064342298e-11,
 9.124121360260552e-11,
 8.825076063406994e-11,
 8.535832048894989e-11,
 8.256068077311755e-11,
 7.985473437944489e-11,
 7.723747603699543e-11,
 7.470599897331893e-11,
 7.225749168613837e-11,
 6.988923482084852e-11,
 6.759859815035253e-11,
 6.538303765388658e-11,
 6.324009269158552e-11,
 6.116738327165496e-11,
 5.916260740711023e-11,
 5.722353855915019e-11,
 5.53480231643239e-11,
 5.353397824274691e-11,
 5.1779389084706924e-11,
 5.0082307013092715e-11,
 4.8440847219158796e-11,
 4.685318666922531e-11,
 4.531756207998447e-11,
 4.3832267960167973e-11,
 4.239565471639819e-11,
 4.100612682112212e-11,
 3.9662141040590415e-11,
 3.8362204720914764e-11,
 3.7104874130301564e-11,
 3.5888752855617746e-11,
 3.471249025151045e-11,
 3.357477994035642e-11,
 3.247435836137711e-11,
 3.141000336730571e-11,
 3.038053286704946e-11,
 2.9384803512838344e-11,
 2.8421709430403943e-11,
 2.7490180990775868e-11,
 2.658918362233348e-11,
 2.5717716661792372e-11,
 2.4874812242850954e-11,
 2.4059534221260996e-11,
 2.3270977135129788e-11,
 2.250826519929814e-11,
 2.177055133267871e-11,
 2.1057016217472765e-11,
 2.0366867389221776e-11,
 1.9699338356682367e-11,
 1.9053687750548326e-11,
 1.842919850007271e-11,
 1.7825177036676758e-11,
 1.7240952523660416e-11,
 1.6675876111159815e-11,
 1.6129320215523138e-11,
 1.560067782230546e-11,
 1.5089361812107756e-11,
 1.4594804308512292e-11,
 1.4116456047388985e-11,
 1.3653785766873268e-11,
 1.3206279617337255e-11,
 1.277344059069974e-11,
 1.2354787968440232e-11,
 1.19498567877047e-11,
 1.155819732490959e-11,
 1.1179374596271475e-11,
 1.0812967874705979e-11,
 1.0458570222561264e-11,
 1.0115788039666234e-11,
 9.78424062619091e-12,
 9.463559759834847e-12,
 9.153389286873265e-12,
 8.853384726606133e-12,
 8.563212888772166e-12,
 8.282551503502112e-12,
 8.011088863399832e-12,
 7.748523477355144e-12,
 7.494563735702168e-12,
 7.248927586352747e-12,
 7.01134222154468e-12,
 6.781543774856371e-12,
 6.559277028152377e-12,
 6.3442951281338605e-12,
 6.136359312178713e-12,
 5.935238643168016e-12,
 5.740709753002967e-12,
 5.552556594528534e-12,
 5.3705702015878575e-12,
 5.1945484569405815e-12,
 5.024295867788064e-12,
 4.8596233486557965e-12,
 4.700348011391552e-12,
 4.546292962046906e-12,
 4.3972871044155265e-12,
 4.253164950010859e-12,
 4.113766434271838e-12,
 3.978936738792232e-12,
 3.848526119376732e-12,
 3.7223897397325335e-12,
 3.6003875106114513e-12,
 3.4823839342245863e-12])
def loop_ana(N0,t):
    for i in range (1,1024):
        N=N0
        N=N*numpy.exp(-t*i/TAU)
        a.append(N)
    return a

error_euler=E-analytic
```

```python
#print(B)
def loop_ana(N0,t):
    for i in range (1,18):
        N=N0
        N=N*numpy.exp(-t*i/TAU)
        a.append(N)
    return a
loop_ana(1000,1)
#print(a)

a=[]
B=[]
C=[]
a.append([1000, 967.2247567687975, 935.5237301064595, 904.8617123036585, 875.2046495923036, 846.5196043248365, 818.7747183931086, 791.9391778462151, 765.983178667987, 740.877893676134, 716.5954405062777, 693.1088506453139, 670.3920394797144, 648.4197773255048, 627.1676614077393, 606.6120887583561, 586.7302300023132, 567.5000040028879])
B.append([1000, 966.6756163192334, 934.4617471861698, 903.3213853879383, 873.2187569542291, 844.1192800602443, 815.9895252991842, 788.7971772786275, 762.5109974966888, 737.1007884553051, 712.5373589694249, 688.7924906322477, 665.8389053979878, 643.6502342449236, 622.2009868827305, 601.4665224692988, 581.4230213033954, 562.0474574606505])
C.append([1000,
 967.2247571093798,
 935.5237307652988,
 904.8617132595269,
 875.2046508250232,
 846.5196058152326,
 818.7747201229663,
 791.9391797982365,
 765.9831808257508,
 740.8778960240569,
 716.5954430295767,
 693.1088533299708,
 670.3920423124417,
 648.4197802937125,
 627.1676644995035,
 606.6120919623893,
 586.7302333079347,
 567.5000073999969])
error_euler=numpy.array(a)-numpy.array(B)
print (error_euler)

error_rk4=abs(numpy.array(a)-numpy.array(C))
print (error_rk4)
```

```python
solve_RK4(f_rad, 1000, 0, 1, 17)
```

Output:
```
[1000,
 967.2247571093798,
 935.5237307652988,
 904.8617132595269,
 875.2046508250232,
 846.5196058152326,
 818.7747201229663,
 791.9391797982365,
 765.9831808257508,
 740.8778960240569,
 716.5954430295767,
 693.1088533299708,
 670.3920423124417,
 648.4197802937125,
 627.1676644995035,
 606.6120919623893,
 586.7302333079347,
 567.5000073999969]
```

```python
solve_euler (f_rad, 1000, 0, 1, 1024)
D=numpy.asarray(solve_euler)
print(D)
```

Output:
```
[1000, 966.6756163192334, 934.4617471861698, 903.3213853879383, 873.2187569542291, 844.1192800602443, 815.9895252991842, 788.7971772786275, 762.5109974966888, 737.1007884553051, 712.5373589694249, 688.7924906322477, 665.8389053979878, 643.6502342449236, 622.2009868827305, 601.4665224692988, 581.4230213033954, 562.0474574606505, 543.3175723414325, 525.211849100224, 507.7094879271232, 490.7903821530742, 474.4350951513751, 458.62483800892966, 443.3414479415907, 428.5673674287985, 414.2856240435451, 400.47981095449217, 387.1340680778438, 374.23306385732167, 361.76197765131144, 349.70648270694625, 338.0527297015686, 326.78733083266303, 315.89734443798176, 305.3702601281952, 295.1939844149877, 285.35682681808845, 275.8474864352764, 266.65503895993214, 257.7689241312216, 249.17893360249434, 240.87519921396054, 232.84818165617344, 225.08865951129428, 217.5877186595505, 210.33674203871692, 203.32739974485628, 196.55163946294607, 190.00167721639914, 183.6699884248507, 177.549299259939, 171.63257
```