# <center>L2 Computational Physics</center>

## <center>Week 3: Differential Equations I</center>

```python
# usual packages to import
import numpy 
import matplotlib.pyplot as plt
%matplotlib inline
```

In this notebook, you will generate and plot the decay curve for Iodine-133 analytically and numerically. $^{133}\textrm{I}$ has a half life $t_{1/2}$ of 20.8 hours. This means that half of the nuclei will have decayed after time $t_{1/2}$. Derive the mean lifetime $\tau$ from that information.

<!-- grade_id: constants, points: 0, grade: false, solution: true -->
```python
# define a function to calculate the mean lifetime from the half life
T_HALF = 20.8
def meanLifetime(halfLife):
    lmb=numpy.log(2)/T_HALF
    TAU=1/lmb
    return TAU
TAU = meanLifetime(T_HALF)

```

Check your average lifetime:

<!-- grade_id: constants_check, points: 1, grade: true, solution: false -->
```python
# this test is worth 1 mark
assert numpy.isclose(TAU, 30.0080568505)         
```

### The Decay Equation

Implement the function `f_rad` such that the differential equation 

$$ \frac{dN}{dt} = f_{rad}(N,t)$$

describes the radioactive decay process.

- *Your function should return values using hours as the time unit.*
- *The function should use the constant* `TAU`.

<!-- grade_id: decay_function, points: 0, grade: false, solution: true -->
```python
def f_rad(N, t):
    # YOUR CODE HERE
    return N*(-1/TAU)


```

Make sure your function works:

<!-- grade_id: decay_test, points: 1, grade: true, solution: false -->
```python
# this test cell is worth 1 mark
assert numpy.isclose(f_rad(1000, 0), -33.324383681)           
```

Solve this first order, ordinary differential equation analytically. Implement this function below, naming it `analytic`. The function should take an initial number of atoms `N0` at time `t=0`, and a time argument. The function should return nuclei count at the time argument. Make sure the function also works for numpy arrays.

<!-- grade_id: analytic_function, points: 0, grade: false, solution: true -->
```python
def analytic(N0, t):
    # YOUR CODE HERE
    N=N0*numpy.exp(-t/TAU)
    return N

```

Check your answer for a single time:

<!-- grade_id: analytic_test, points: 1, grade: true, solution: false -->
```python
# this test is worth 1 mark
assert numpy.isclose(analytic(1000, 41.6), 250.0)           
```

<!-- grade_id: analytic_test_array, points: 2, grade: true, solution: false -->
```python
# this test is worth 1 mark
assert numpy.isclose(analytic(1000, numpy.arange(0, 60, 6)), 
                     [1000.        ,  818.77471839,  670.39203948,  548.90005334,
                       449.4254866 ,  367.97822623,  301.29126855,  246.68967356,
                       201.983268  ,  165.37879338]).all()

```

## Numerically Solving the ODE

We now wish to solve our differential equation numerically. We shall do this using Euler's and RK4 methods.

### Euler's Method

Create a function which takes as its arguments the initial number of atoms, `n0`, the initial time `t0`, the time step, `dt`, and the number of steps to perform, `n_steps`.  This function should return an array of the number of counts at each time step using Euler's method. This array should contain the initial and final values, so the array length should be `n_steps+1` 

<!-- grade_id: euler_function, points: 0, grade: false, solution: true -->
```python

def solve_euler(f, n0, t0, dt, n_panels):
    B=[n0]
    n=n0
    # YOUR CODE HERE
    #x(t)+dt*f(x(t),t)
    for i in range (0,n_panels):
        n=n+dt*f(n,t0+i*dt)
        B.append(n)
    print (B)
    return B

```

Try your solution:

<!-- grade_id: euler_test_length, points: 1, grade: true, solution: false -->
```python
# this test is worth 1 mark
assert len(solve_euler(f_rad, 1000, 0, 1, 17)) == 18
```

Output:
```
[1000, 966.6756163192334, 934.4617471861698, 903.3213853879383, 873.2187569542291, 844.1192800602443, 815.9895252991842, 788.7971772786275, 762.5109974966888, 737.1007884553051, 712.5373589694249, 688.7924906322477, 665.8389053979878, 643.6502342449236, 622.2009868827305, 601.4665224692988, 581.4230213033954, 562.0474574606505]
```

<!-- grade_id: euler_test_one, points: 2, grade: true, solution: false -->
```python
# this test is worth 2 marks
assert numpy.isclose(solve_euler(f_rad, 1000, 0, 6, 1), [1000.,  800.05369792]).all()
```

Output:
```
[1000, 800.0536979154003]
```

<!-- grade_id: euler_test_many, points: 2, grade: true, solution: false -->
```python
# this test is worth 2 mark
assert numpy.isclose(solve_euler(f_rad, 1000, 0, 6, 10), [1000.        ,  800.05369792,  640.08591955,  512.10310692,
                                                409.7099844 ,  327.7899881 ,  262.24959212,  209.81375595,
                                                167.86227132,  134.29883091,  107.4462763 ]).all()
```

Output:
```
[1000, 800.0536979154003, 640.0859195481066, 512.1031069180422, 409.7099844037453, 327.78998809507743, 262.2495921151118, 209.8137559485006, 167.86227132011726, 134.29883091013807, 107.44627629537104]
```

### RK 4 method

Implement the RK4 method in the `solve_RK4` function. The arguments are the same as for `solve_euler`.

<!-- grade_id: RK4, points: 0, grade: false, solution: true -->
```python
def solve_RK4(f, n0, t0, dt, nsteps):
    n=n0
    b=[n0]
    for i in range (1,nsteps+1):
        k1=dt*f( n,t0)
        k2=dt*f(n+0.5*k1,t0+0.5*dt)
        k3=dt*f(n+0.5*k2,t0+0.5*dt) 
        k4=dt*f(n+k3,t0+dt)
        k=( k1+2*k2+2*k3+ k4)/6 
        n=n+k
        t0=t0+dt
        b.append(n)
    return b
  
    

```

<!-- grade_id: RK4-test-1, points: 1, grade: true, solution: false -->
```python
# This checks that we return an array of the right length
# this test is worth 1 mark
assert len(solve_RK4(f_rad, 1000, 0, 1, 17)) == 18
```

<!-- grade_id: RK4-test-2, points: 2, grade: true, solution: false -->
```python
# This checks that a single step is working
# this test is worth 2 mark
assert numpy.isclose(solve_RK4(f_rad, 1000,0, 6, 1), [1000.,  818.7773]).all()
```

<!-- grade_id: RK4-test-3, points: 2, grade: true, solution: false -->
```python
# This checks multiple steps
# this test is worth 2 marks
assert numpy.isclose(solve_RK4(f_rad, 1000, 0, 6, 10), [
    1000.,
    818.77729521,  
    670.39625915,  
    548.90523578,
    449.43114428,  
    367.9840167,  
    301.29695787,  
    246.69510822, 
    201.98835345,  
    165.3834777,  
    135.41223655]).all()
```

## Plotting task

**Task 1: **

Create a plot to show that the RK4 method has an error that scales better with the number of steps than the Euler method. (click on the "+" button to create new cells.)       [task worth 5 marks]


```python
#finding the error in euler
panel_counts =[4, 8, 16, 32, 64, 128, 256, 512, 1024]
from numpy import matrix
N0=1000


A=[N0]
def loop_ana(N0,t0,t,npanels):
    for i in range (t0,npanels-t+1):
        N=N0
        N=N*numpy.exp(-(i+t)/TAU)
        A.append(N)
    return A
#euler4=matrix(loop_ana(1000,1,0,5))-matrix(solve_euler (f_rad,1000,0,1,5))
#print(euler4)

C=[]
def loop_anav2(N0,t0,t,npanels):
    for i in range (t0,npanels-t+1):
        N=N0
        N=N*numpy.exp(-(i+t)/TAU)
        C.append(N)
    return C

def solve_eulerv2(f, n0, t0, dt, n_panels):
    B=[]
    n=n0
    for i in range (0,n_panels):
        n=n+dt*f(n,t0+i*dt)
        B.append(n)
    print (B)
    return B

#error_euler64=abs(matrix(loop_anav2(1000,5,1,69))-matrix(solve_eulerv2 (f_rad,1000,4,1,64)))
print(loop_anav2(1000,5,1,69))
#solve_euler (f_rad,1000,4,1,64)
print (solve_eulerv2(f_rad,1000,4,1,64))
```

Output:
```
[818.7747183931086, 791.9391778462151, 765.983178667987, 740.877893676134, 716.5954405062777, 693.1088506453139, 670.3920394797144, 648.4197773255048, 627.1676614077393, 606.6120887583561, 586.7302300023132, 567.5000040028879, 548.9000533379849, 530.9097205802124, 513.5090253543862, 496.67864214697846, 480.39987884286785, 464.6546559655526, 449.42548659777094, 434.69545696022743, 420.44820762685725, 406.66791535576385, 393.3392755156526, 380.4474850882421, 367.9782262277757, 355.9176503593739, 344.25236279856733, 332.96940787492815, 322.0562545432779, 311.5007824664919, 301.29126855444264, 291.4163739441333, 281.8651314065592, 272.62693316631436, 263.6915191204116, 255.04896544323492, 246.68967356496628, 238.6043595112486, 230.7840435922421, 223.22004042962593, 215.9039493104661, 208.8276448572383, 201.98326800364316, 195.3632172661906, 188.96014030186092, 182.76692574246528, 176.7766952966369, 170.98279611068142, 165.37879337980274, 159.95846320149693, 154.7157856631785, 149.644938156
```

```python
A=[]
def loop_ana(N0,t0,t,npanels):
    for i in range (t0,npanels-t+1):
        N=N0
        N=N*numpy.exp(-(t+i)/TAU)
        A.append(N)
    return A
print(loop_ana (1000,1,0,5))
#loop_ana (1000,4,1,8)

print(solve_euler(f_rad, 1000, 0, 1, 6))
```

Output:
```
[967.2247567687975, 935.5237301064595, 904.8617123036585, 875.2046495923036, 846.5196043248365]
[1000, 966.6756163192334, 934.4617471861698, 903.3213853879383, 873.2187569542291, 844.1192800602443, 815.9895252991842]
[1000, 966.6756163192334, 934.4617471861698, 903.3213853879383, 873.2187569542291, 844.1192800602443, 815.9895252991842]
```

```python
solve_RK4(f_rad, 1000, 0, 1, 17)
```

Output:
```
[1000,
 967.2247571093798,
 935.5237307652988,
 904.8617132595269,
 875.2046508250232,
 846.5196058152326,
 818.7747201229663,
 791.9391797982365,
 765.9831808257508,
 740.8778960240569,
 716.5954430295767,
 693.1088533299708,
 670.3920423124417,
 648.4197802937125,
 627.1676644995035,
 606.6120919623893,
 586.7302333079347,
 567.5000073999969]
```

```python
solve_euler (f_rad, 1000, 0, 1, 1024)
D=numpy.asarray(solve_euler)
print(D)
```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-1-006e3d937fcb> in <module>()
----> 1 solve_euler (f_rad, 1000, 0, 1, 1024)
      2 D=numpy.asarray(solve_euler)
      3 print(D)

NameError: name 'solve_euler' is not defined
```