# <center>L2 Computational Physics</center>

## <center>Week 3: Differential Equations I</center>

```python
# usual packages to import
import numpy 
import matplotlib.pyplot as plt
%matplotlib inline
```

In this notebook, you will generate and plot the decay curve for Iodine-133 analytically and numerically. $^{133}\textrm{I}$ has a half life $t_{1/2}$ of 20.8 hours. This means that half of the nuclei will have decayed after time $t_{1/2}$. Derive the mean lifetime $\tau$ from that information.

<!-- grade_id: constants, points: 0, grade: false, solution: true -->
```python
# define a function to calculate the mean lifetime from the half life
def meanLifetime(halfLife):
    # YOUR CODE HERE
    tau = (halfLife)/numpy.log(2);
    return tau;

T_HALF = 20.8
TAU = meanLifetime(T_HALF)

```

Check your average lifetime:

<!-- grade_id: constants_check, points: 1, grade: true, solution: false -->
```python
# this test is worth 1 mark
assert numpy.isclose(TAU, 30.0080568505)         
```

### The Decay Equation

Implement the function `f_rad` such that the differential equation 

$$ \frac{dN}{dt} = f_{rad}(N,t)$$

describes the radioactive decay process.

- *Your function should return values using hours as the time unit.*
- *The function should use the constant* `TAU`.

<!-- grade_id: decay_function, points: 0, grade: false, solution: true -->
```python
def f_rad(N, t):
    # YOUR CODE HERE
    r = -(1/TAU)*N;
    #r = (-(N)*numpy.exp(-(t/TAU)))/TAU;
    return r;


#see time dependence
```

```python
f_rad(1000,6)
```

Output:
```
-27.285162863843578
```

Make sure your function works:

<!-- grade_id: decay_test, points: 1, grade: true, solution: false -->
```python
# this test cell is worth 1 mark
assert numpy.isclose(f_rad(1000, 0), -33.324383681)           
```

Solve this first order, ordinary differential equation analytically. Implement this function below, naming it `analytic`. The function should take an initial number of atoms `N0` at time `t=0`, and a time argument. The function should return nuclei count at the time argument. Make sure the function also works for numpy arrays.

<!-- grade_id: analytic_function, points: 0, grade: false, solution: true -->
```python
def analytic(N0, t):
    # YOUR CODE HERE
    Nt = (N0)*numpy.exp(-(t/TAU));
    return Nt;
```

Check your answer for a single time:

<!-- grade_id: analytic_test, points: 1, grade: true, solution: false -->
```python
# this test is worth 1 mark
assert numpy.isclose(analytic(1000, 41.6), 250.0)           
```

<!-- grade_id: analytic_test_array, points: 2, grade: true, solution: false -->
```python
# this test is worth 1 mark
assert numpy.isclose(analytic(1000, numpy.arange(0, 60, 6)), 
                     [1000.        ,  818.77471839,  670.39203948,  548.90005334,
                       449.4254866 ,  367.97822623,  301.29126855,  246.68967356,
                       201.983268  ,  165.37879338]).all()

```

## Numerically Solving the ODE

We now wish to solve our differential equation numerically. We shall do this using Euler's and RK4 methods.

### Euler's Method

Create a function which takes as its arguments the initial number of atoms, `n0`, the initial time `t0`, the time step, `dt`, and the number of steps to perform, `n_steps`.  This function should return an array of the number of counts at each time step using Euler's method. This array should contain the initial and final values, so the array length should be `n_steps+1` 

<!-- grade_id: euler_function, points: 0, grade: false, solution: true -->
```python
def solve_euler(f, n0, t0, dt, n_panels):
    # YOUR CODE HERE
    ar = [n0]
    nx = n0
    for i in range(1,n_panels+1):
        #nx = nx + ((dt)*f(nx,t0+(dt*(i-1))));
        nx = nx + ((dt)*f(nx,t0));
        ar.append(nx);
    return ar;
```

```python
print(solve_euler(f_rad,1000,20,6,10))
```

Output:
```
[1000, 800.0536979154003, 640.0859195481066, 512.1031069180422, 409.7099844037453, 327.78998809507743, 262.2495921151118, 209.8137559485006, 167.86227132011726, 134.29883091013807, 107.44627629537104]
```

Try your solution:

<!-- grade_id: euler_test_length, points: 1, grade: true, solution: false -->
```python
# this test is worth 1 mark
assert len(solve_euler(f_rad, 1000, 0, 1, 17)) == 18
```

<!-- grade_id: euler_test_one, points: 2, grade: true, solution: false -->
```python
# this test is worth 2 marks
assert numpy.isclose(solve_euler(f_rad, 1000, 0, 6, 1), [1000.,  800.05369792]).all()
```

<!-- grade_id: euler_test_many, points: 2, grade: true, solution: false -->
```python
# this test is worth 2 mark
assert numpy.isclose(solve_euler(f_rad, 1000, 0, 6, 10), [1000.        ,  800.05369792,  640.08591955,  512.10310692,
                                                409.7099844 ,  327.7899881 ,  262.24959212,  209.81375595,
                                                167.86227132,  134.29883091,  107.4462763 ]).all()
```

### RK 4 method

Implement the RK4 method in the `solve_RK4` function. The arguments are the same as for `solve_euler`.

<!-- grade_id: RK4, points: 0, grade: false, solution: true -->
```python
def solve_RK4(f, n0, t0, dt, n_steps):
    # YOUR CODE HERE
    ar = [n0]
    nx = n0
for i in range(1,n_steps+1):
    k1 = f(nx,t0);
    k2 = f(nx+(k1*(dt/2)),t0);
    k3 = f(nx+(k2*(dt/2)),t0);
    k4 = f(nx+(k3*dt),t0);
    k = (1/6)*(k1+(2*k2)+(2*k3)+k4);
    nx = n0 +(k*dt);
    ar.append(nx);
return ar;
```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-248-8cd28260cfeb> in <module>()
      3     ar = [n0]
      4     nx = n0
----> 5 for i in range(1,n_steps+1):
      6     k1 = f(nx,t0);
      7     k2 = f(nx+(k1*(dt/2)),t0);

NameError: name 'n_steps' is not defined
```

<!-- grade_id: RK4-test-1, points: 1, grade: true, solution: false -->
```python
# This checks that we return an array of the right length
# this test is worth 1 mark
assert len(solve_RK4(f_rad, 1000, 0, 1, 17)) == 18
```

Output:
```
TypeError                                 Traceback (most recent call last)
<ipython-input-230-30aafc191006> in <module>()
      1 # This checks that we return an array of the right length
      2 # this test is worth 1 mark
----> 3 assert len(solve_RK4(f_rad, 1000, 0, 1, 17)) == 18

TypeError: object of type 'NoneType' has no len()
```

<!-- grade_id: RK4-test-2, points: 2, grade: true, solution: false -->
```python
# This checks that a single step is working
# this test is worth 2 mark
assert numpy.isclose(solve_RK4(f_rad, 1000,0, 6, 1), [1000.,  818.7773]).all()
```

Output:
```
<ipython-input-231-42020d739665> in <module>()
      1 # This checks that a single step is working
      2 # this test is worth 2 mark
----> 3 assert numpy.isclose(solve_RK4(f_rad, 1000,0, 6, 1), [1000.,  818.7773]).all()

/opt/conda/lib/python3.6/site-packages/numpy/core/numeric.py in isclose(a, b, rtol, atol, equal_nan)
   2361     y = array(y, dtype=dt, copy=False, subok=True)
   2362 
-> 2363     xfin = isfinite(x)
   2364     yfin = isfinite(y)
   2365     if all(xfin) and all(yfin):

TypeError: ufunc 'isfinite' not supported for the input types, and the inputs could not be safely coerced to any supported types according to the casting rule ''safe''
```

<!-- grade_id: RK4-test-3, points: 2, grade: true, solution: false -->
```python
# This checks multiple steps
# this test is worth 2 marks
assert numpy.isclose(solve_RK4(f_rad, 1000, 0, 6, 10), [
    1000.,
    818.77729521,  
    670.39625915,  
    548.90523578,
    449.43114428,  
    367.9840167,  
    301.29695787,  
    246.69510822, 
    201.98835345,  
    165.3834777,  
    135.41223655]).all()
```

Output:
```
<ipython-input-232-0187a5525018> in <module>()
     12     201.98835345,
     13     165.3834777,
---> 14     135.41223655]).all()

/opt/conda/lib/python3.6/site-packages/numpy/core/numeric.py in isclose(a, b, rtol, atol, equal_nan)
   2361     y = array(y, dtype=dt, copy=False, subok=True)
   2362 
-> 2363     xfin = isfinite(x)
   2364     yfin = isfinite(y)
   2365     if all(xfin) and all(yfin):

TypeError: ufunc 'isfinite' not supported for the input types, and the inputs could not be safely coerced to any supported types according to the casting rule ''safe''
```

## Plotting task

**Task 1: **

Create a plot to show that the RK4 method has an error that scales better with the number of steps than the Euler method. (click on the "+" button to create new cells.)       [task worth 5 marks]


```python
n_steps = [1, 4, 8, 16, 32, 64, 128, 256, 512, 1024]
n0 = 1000
t0 = 0
dt = 1
errE = []
errRK4 = []
for i in n_steps:
    errE.append(numpy.abs((solve_euler(f_rad, n0, t0, dt, i)-analytic(n0, t0+(dt*(1-i))))));
    #errRK4.append(numpy.abs((solve_RK4(f_rad, n0, t0, dt, i)-analytic(n0, t0+(dt*(1-i))))));

print(errE)
    

```

Output:
```
[array([ 0.        , 33.32438368]), array([105.14124579, 138.46562947, 170.6794986 , 201.8198604 ,
       231.92248883]), array([262.72323428, 296.04761796, 328.26148709, 359.40184889,
       389.50447732, 418.60395422, 446.73370898, 473.926057  ,
       500.21223678]), array([ 648.49995332,  681.824337  ,  714.03820613,  745.17856793,
        775.28119636,  804.38067326,  832.51042802,  859.70277604,
        885.98895582,  911.39916486,  935.96259435,  959.70746269,
        982.66104792, 1004.84971907, 1026.29896643, 1047.03343085,
       1067.07693201]), array([1809.63868746, 1842.96307114, 1875.17694028, 1906.31730207,
       1936.41993051, 1965.5194074 , 1993.64916216, 2020.84151018,
       2047.12768996, 2072.53789901, 2097.10132849, 2120.84619683,
       2143.79978206, 2165.98845322, 2187.43770058, 2208.17216499,
       2228.21566616, 2247.59123   , 2266.32111512, 2284.42683836,
       2301.92919953, 2318.84830531, 2335.20359231, 2351.01384945,
       2366.29723952, 2381.07132003
```

```python

```