# Assignment 1

First we need to import a few modules, execute the cell below:

```python
import numpy
import matplotlib.pyplot as plt
%matplotlib inline
```

Define a function `f` that is equivalent to $f(x)=\cos(x)$. It should work both for single arguments and `numpy` arrays.

<!-- grade_id: f_definition, points: 0, grade: false, solution: true -->
```python
def f(x):
    import numpy
    '''Function equivalent to cos(x), should work for one argument or a numpy array'''
    return np.cos(x)
```

You can use the following cell to check that your code works for scalar arguments:

<!-- grade_id: f_definition_test, points: 1, grade: true, solution: false -->
```python
assert numpy.isclose(f(0.5),0.87758256189)
```

Output:
```
<ipython-input-13-5c26e2f4934c> in <module>()
----> 1 assert numpy.isclose(f(0.5),0.87758256189)

<ipython-input-12-226df496802b> in f(x)
      2     import numpy
      3     '''Function equivalent to cos(x), should work for one argument or a numpy array'''
----> 4     return np.cos(x)

NameError: name 'np' is not defined
```

Make sure your code also works with a numpy array:

<!-- grade_id: f_definition_nparray_test, points: 2, grade: true, solution: false -->
```python
xs = numpy.arange(start=0.1, stop=0.7, step=0.1)
assert numpy.isclose(
            f(xs), 
            [0.99500417, 0.98006658, 0.95533649, 0.92106099, 0.87758256, 0.82533561]
        ).all()
```

Output:
```
<ipython-input-14-d78de3b75973> in <module>()
      1 xs = numpy.arange(start=0.1, stop=0.7, step=0.1)
      2 assert numpy.isclose(
----> 3             f(xs),
      4             [0.99500417, 0.98006658, 0.95533649, 0.92106099, 0.87758256, 0.82533561]
      5         ).all()

<ipython-input-12-226df496802b> in f(x)
      2     import numpy
      3     '''Function equivalent to cos(x), should work for one argument or a numpy array'''
----> 4     return np.cos(x)

NameError: name 'np' is not defined
```

Define the analytic derivative of the function`f` here (again, it should work both for single numbers and numpy arrays).

<!-- grade_id: analytic_derivative, points: 0, grade: false, solution: true -->
```python
def df_analytic(x):
    '''
    The analytic derivative
    '''
    # YOUR CODE HERE
```

<!-- grade_id: g_definition_test, points: 1, grade: true, solution: false -->
```python
assert numpy.isclose(df_analytic(0.5), -0.4794255386)
```

Output:
```
<ipython-input-16-98d4b7aab4b5> in <module>()
----> 1 assert numpy.isclose(df_analytic(0.5), -0.4794255386)

/opt/conda/lib/python3.6/site-packages/numpy/core/numeric.py in isclose(a, b, rtol, atol, equal_nan)
   2361     y = array(y, dtype=dt, copy=False, subok=True)
   2362 
-> 2363     xfin = isfinite(x)
   2364     yfin = isfinite(y)
   2365     if all(xfin) and all(yfin):

TypeError: ufunc 'isfinite' not supported for the input types, and the inputs could not be safely coerced to any supported types according to the casting rule ''safe''
```

<!-- grade_id: g_definition_nparray_test, points: 2, grade: true, solution: false -->
```python
xs = numpy.arange(start=0.1, stop=0.7, step=0.1)
assert numpy.isclose(
            df_analytic(xs),
            [-0.09983342, -0.19866933, -0.29552021, -0.38941834, -0.47942554,-0.56464247]
        ).all()
```

Output:
```
<ipython-input-17-bb02cfb57d67> in <module>()
      2 assert numpy.isclose(
      3             df_analytic(xs),
----> 4             [-0.09983342, -0.19866933, -0.29552021, -0.38941834, -0.47942554,-0.56464247]
      5         ).all()

/opt/conda/lib/python3.6/site-packages/numpy/core/numeric.py in isclose(a, b, rtol, atol, equal_nan)
   2361     y = array(y, dtype=dt, copy=False, subok=True)
   2362 
-> 2363     xfin = isfinite(x)
   2364     yfin = isfinite(y)
   2365     if all(xfin) and all(yfin):

TypeError: ufunc 'isfinite' not supported for the input types, and the inputs could not be safely coerced to any supported types according to the casting rule ''safe''
```

Now define the numerical derivative using the backward difference method. The function `backward_difference` takes three arguments, `f`, the function to calculate the derivative for, `x` the position to calculate the derivative at and `dx` the interval length.

<!-- grade_id: backward_difference, points: 0, grade: false, solution: true -->
```python
def backward_difference(f, x, dx):
    '''
    This function implements the backward difference method for the 
    first derivative of the function f at position x using interval
    dx.
    '''
    # YOUR CODE HERE
```

<!-- grade_id: backward_difference_test, points: 2, grade: true, solution: false -->
```python
assert numpy.isclose(
    backward_difference(f, 0.1, 0.0001),
    -0.09978366627194646
)
```

Output:
```
<ipython-input-19-f557846ec4dc> in <module>()
      1 assert numpy.isclose(
      2     backward_difference(f, 0.1, 0.0001),
----> 3     -0.09978366627194646
      4 )

/opt/conda/lib/python3.6/site-packages/numpy/core/numeric.py in isclose(a, b, rtol, atol, equal_nan)
   2361     y = array(y, dtype=dt, copy=False, subok=True)
   2362 
-> 2363     xfin = isfinite(x)
   2364     yfin = isfinite(y)
   2365     if all(xfin) and all(yfin):

TypeError: ufunc 'isfinite' not supported for the input types, and the inputs could not be safely coerced to any supported types according to the casting rule ''safe''
```

**Task 1:**

Here is a skeleton code to plot the difference between the analytical derivative and the numerical implementation. Modify and expand it to provide a plot with three curves for the difference for a case where $dx$ is too large, about right and too small. The plot should have a legend, labelled axes and a title.

You can get help about plotting on the Python handbook that you can find at:

https://dmaitre.phyip3.dur.ac.uk/notes/l1python/part-5/


```python
xs = numpy.linspace(-2*numpy.pi,2*numpy.pi,100)
df_dx_1 = backward_difference(f, xs, dx=1e-4)
df_dx_2 = backward_difference(f, xs, dx=1e-6)
df_dx_analytical = df_analytic(xs)
plt.figure(figsize=(8, 4))
plt.plot(xs, df_dx_1 - df_dx_analytical)
plt.plot(xs, df_dx_2 - df_dx_analytical)
```

Output:
```
TypeError                                 Traceback (most recent call last)
<ipython-input-20-2d0e8f23f5d5> in <module>()
      4 df_dx_analytical = df_analytic(xs)
      5 plt.figure(figsize=(8, 4))
----> 6 plt.plot(xs, df_dx_1 - df_dx_analytical)
      7 plt.plot(xs, df_dx_2 - df_dx_analytical)

TypeError: unsupported operand type(s) for -: 'NoneType' and 'NoneType'
<Figure size 576x288 with 0 Axes>
```

**Task 2: ** 

Describe in the text box below why the numerical method looses accuracy in the case where
$\Delta x$ is too small and the case where $\Delta x$ is too large. (Double-click on "YOUR ANSWER HERE" to edit.)


<!-- grade_id: textAnswer, points: 4, grade: true, solution: true -->
YOUR ANSWER HERE

```python

```

```python

```

```python

```

```python

```

```python

```

```python

```