# <center>Computational Physics</center>
---

## Week 2: Numerical Integration

```python
import numpy 
import matplotlib.pyplot as plt
%matplotlib inline
```

Define the function `f`, such that $\textrm{f}(x) \equiv x^{2}\sin(x)$. This is the function that we will be integrating.

<!-- grade_id: integrand, points: 0, grade: false, solution: true -->
```python
def f(x):
    '''Function equivalent to x^2 sin(x).'''
    # YOUR CODE HERE
    return (numpy.sin(x))*x**2
```

Ensure your function works with numpy arrays:

<!-- grade_id: cell-653afd49e7782cc4, points: 1, grade: true, solution: false -->
```python
xs=numpy.arange(0, 1, step=0.1)
assert numpy.isclose(f(xs), 
                     [0., 0.00099833, 0.00794677, 0.02659682, 0.06230693,
                      0.11985638, 0.20327129, 0.31566667, 0.4591079 , 0.6344948 ]).all()
```

Derive the indefinite integral of $f(x)$ nalytically. Call this function $g(x)$ and implement it below. Set the constant of integration such that $g(0)=0$.

<!-- grade_id: integral, points: 0, grade: false, solution: true -->
```python
def g(x):
    '''Analytical integral of f(x).'''
    # YOUR CODE HERE
    return (-x**2*numpy.cos(x)+2*(x*numpy.sin(x)+numpy.cos(x))-2)
```

Check your solution with the same numpy array:

<!-- grade_id: cell-648a123ee323b45e, points: 1, grade: true, solution: false -->
```python
assert g(0) == 0.
```

<!-- grade_id: cell-f8eaec03e0540a78, points: 1, grade: true, solution: false -->
```python
assert numpy.isclose(g(xs), 
                     [0., 0.00002497, 0.00039822, 0.00200482, 0.0062869, 
                      0.01519502, 0.03112138, 0.05681646, 0.09529087, 0.1497043 ]).all()
```

Now, using the analytically derived indefinite integral, $g(x)$, define a function which calculates the definite integral of $f(x)$ over the interval $(x_{min},~x_{max})$.

<!-- grade_id: cell-6f8d8cc5f5c7564c, points: 0, grade: false, solution: true -->
```python
def integrate_analytic(xmin, xmax):
    '''Analytical integral of f(x) from xmin to xmax.'''
    # YOUR CODE HERE
    return g(xmax)-g(xmin)
```

Check your analytic function:

<!-- grade_id: analytical_integral, points: 1, grade: true, solution: false -->
```python
assert numpy.isclose(integrate_analytic(xmin=0, xmax=4), 1.096591)
```

## Numerical implementation

Create a function which calculates the definite integral of the function $f(x)$ over the interval $(x_{min},~x_{max})$ using Simpson's rule with $N$ panels.

<!-- grade_id: numerical_integral, points: 0, grade: false, solution: true -->
```python
def integrate_numeric(xmin, xmax, N):
    ''' 
    Numerical integral of f from xmin to xmax using Simpson's rule with 
        N panels.
    '''
    if N == 1:
        deltax = (xmax-xmin)/N
        m1 =  xmin + deltax/2
            
        return ((deltax/6)*(f(xmin)+f(xmax)+4*f(m1)))
    else:


        deltax = (xmax-xmin)/N

        m1 =  xmin + deltax/2

        x1 = xmin + deltax
    
        fx=0
        fm=0
    
     
        for i in range(1, (N)):
            fm = fm + f(m1 + deltax*i)
            print(fm)
    
        for i in range (1, (N-1)):
            fx = fx+ f(x1 + deltax*i)

        return ((deltax/6)*(f(xmin)+ f(xmax) + 4*f(m1)+ 4*fm + 2*f(x1)+2*fx))

```

Output:
```
2.2443637198591224
2.4628353491598687
```

Make sure you have implemented Simpson's rule correctly:

<!-- grade_id: cell-038d65c4eff18248, points: 3, grade: true, solution: false -->
```python
assert numpy.isclose(integrate_numeric(xmin=0, xmax=4, N=1), 1.6266126)
```

<!-- grade_id: cell-3304084098accd44, points: 3, grade: true, solution: false -->
```python
assert numpy.isclose(integrate_numeric(xmin=0, xmax=4, N=50), 1.096591)
```

Output:
```
0.0017238557849604387
0.009670629016762888
0.031336911864189396
0.07699165249664106
0.15945353293341857
0.29380992220636504
0.49708121262857785
0.7878351069347178
1.1857561336305722
1.7111763190820106
2.384573530280791
3.226044515088687
4.25476010103929
5.48841036420258
6.942647841756191
8.630537032202257
10.56201850297964
12.743395904854887
15.176854075206533
17.860016198402583
20.78554768215089
23.940814006196245
27.307599307229772
30.861891885489598
34.57374215929718
38.40719785958843
42.320320454075244
46.26528592748295
50.188572128516526
54.03123393460004
57.72926649133777
61.21405576485087
64.4129146118084
67.2497015345327
69.64551825763309
71.51948124890933
72.78956132144813
73.37348450647526
73.18968648796456
72.1583120502659
70.2022502187
67.24819507926983
63.22772165584146
58.078365709143974
51.74469590874914
44.17936652298551
35.34413857777265
25.210857357899734
13.762374166562175
```

## Plotting task

** Task 1 **

There will always be some discrepancy between a numerically calculated result and an analytically derived result. Produce a log-log plot showing the fractional error between these two results as the number of panels is varied. The plot should have labels and a title.


```python
x0, x1 = 0, 2  # Bounds to integrate f(x) over
panel_counts = [4, 8, 16, 32, 64, 128, 256, 512, 1024]  # Panel numbers to use
result_analytic = integrate_analytic(x0, x1)  # Define reference value from analytical solution

plt.plot(panel_counts, integrate_numeric(x0, x1, panel_counts))


plt.loglog
plt.ylabel('fractional error between numerical and analytical integration of the function x^2sin(x)')
plt.xlabel('number of panels used for numerical integration')
plt.title('')

```

Output:
```
<ipython-input-106-e137534ac615> in <module>()
      3 result_analytic = integrate_analytic(x0, x1)  # Define reference value from analytical solution
      4 
----> 5 plt.plot(panel_counts, integrate_numeric(x0, x1, panel_counts))
      6 
      7 

<ipython-input-102-97878c9629ea> in integrate_numeric(xmin, xmax, N)
     12 
     13 
---> 14         deltax = (xmax-xmin)/N
     15 
     16         m1 =  xmin + deltax/2

TypeError: unsupported operand type(s) for /: 'int' and 'list'
```

What effect(s) does changing the number of panels used have
on the accuracy of the numerical method? What happens if the number of panels is taken too large?

<!-- grade_id: answer1, points: 3, grade: true, solution: true -->
Increasing the number of panels increases the accuracy of the numerical method as error in numerical integration scales with factor N^-4. If number of panels is taken too high, numerical round-off can cause higher sporadic errors, that deviate from this trend.

If the trapezium rule was being used, how would the panel
count affect accuracy? 

<!-- grade_id: cell-eeb89d1f4669ff1d, points: 1, grade: true, solution: true -->
Error scales as N^-2 which behaves better to much higher N. It is a lower order approximation so will not have numerical round off issues.