# <center>L2 Computational Physics</center>
---

## Week 5: Monte Carlo Methods

In this notebook, you will simulate a system with of three nuclei $A$, $B$ and $C$ where $A$ decays into $B$ and $B$ decays into $C$. If exposed to a neutron flux nucleus $C$ can be activated into a nucleus $A$. 

```python
import numpy
from matplotlib import pyplot as plt
import random
```

Implement a function that tells whether a transition has occured, based on the transition probability and a random number. Use the random number `r` from `random.random()` and use the procedure described in the notes so that the checks can work in a reproducible way.

<!-- grade_id: decay_function, points: 0, grade: false, solution: true -->
```python
def has_transitioned(prob):
    r = random.random()
    if r<prob:
        return True
        print(True)
    else:
        return False
        print(False)
```

<!-- grade_id: decay_function_test, points: 1, grade: true, solution: false -->
```python
# this test is worth 1 mark
random.seed(9867)
assert [ has_transitioned(0.5) for i in range(10)] == [False, False, True, False, False, False, False, True, False, True]
```

We are going to keep track of the state of the atoms using a either `'A'`, `'B'` or `'C'` for each nucleus. For example 
```python
state = ['A', 'A', 'B', 'A', 'C']
```

We will capture the possible physical reaction in a list of allowed transitions such as

`('A', 'B', 0.25 )`

where the first element is the initial state, the second element is the final state of the transition and its probability in the time interval considered is the third argument.

Define a function that takes as arguments the current state and a list of such transition rules and implements the transition (or not) and returns the new state

<!-- grade_id: evolve_function, points: 0, grade: false, solution: true -->
```python
def evolveOne(currentState, rules):
    if len(rules)==2:
        rules.append(('C','A',0.0))
    first_states=[]
    final_states=[]
    probs=[]
    for rule in rules:
        first_states.append(rule[0])
        final_states.append(rule[1])
        probs.append(rule[2])
    index=first_states.index(currentState)
    prob=probs[index]
    if has_transitioned(prob):
        if currentState=='A':
            return 'B'
        elif currentState=='B':
            return 'C'
        else:
            return 'A'
    else:
        return currentState
        
```

<!-- grade_id: cell-7782dc2753905180, points: 1, grade: true, solution: false -->
```python
# these tests are worth 1 mark
alwaysDecayRules = [
    ('A', 'B', 1.0),
    ('B', 'C', 1.0)
]
assert evolveOne('A', alwaysDecayRules) == 'B'
assert evolveOne('B', alwaysDecayRules) == 'C'
```

<!-- grade_id: cell-e2a70f6b9e202d32, points: 2, grade: true, solution: false -->
```python
# these tests are worth 2 mark
random.seed(112211)
testRules = [
    ('A', 'B', 0.33),
    ('B', 'C', 0.75)
]
assert evolveOne('A', testRules) == 'A'
assert evolveOne('A', testRules) == 'A'
assert evolveOne('A', testRules) == 'A'
assert evolveOne('A', testRules) == 'A'
assert evolveOne('A', testRules) == 'B'

assert evolveOne('B', testRules) == 'B'
assert evolveOne('B', testRules) == 'C'
assert evolveOne('B', testRules) == 'C'
assert evolveOne('B', testRules) == 'C'
assert evolveOne('B', testRules) == 'C'

# with no rules there should be no change
assert evolveOne('C', testRules) == 'C'

```

Now implement a function that takes a list of states and transition them according to the rules passed as argument. This function should return a new vector of states, it should not modify the state passed as an argument!

<!-- grade_id: cell-b6743e72d540ebc3, points: 0, grade: false, solution: true -->
```python
def evolveMany(states, rules):
    newState = []
    for state in states:
        newState.append(evolveOne(state, rules))
    #print(newState)
    return newState

```

<!-- grade_id: cell-231f8396b20ac4be, points: 1, grade: true, solution: false -->
```python
# these tests are worth 1 mark
random.seed(112287)
testRules = [
    ('A', 'B', 0.33),
    ('B', 'C', 0.75)
]
initialTestArray = ['A','B','C']*5
evolvedTest = evolveMany(initialTestArray, testRules)
targetArray = ['B', 'C', 'C', 'A', 'C', 'C', 'A', 'B', 'C', 'A', 'C', 'C', 'B', 'C', 'C'] 
assert evolvedTest == targetArray
# checks the initial array is left unchanged
assert initialTestArray == ['A','B','C']*5
```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-165-8c49ec61fa06> in <module>()
      8 evolvedTest = evolveMany(initialTestArray, testRules)
      9 targetArray = ['B', 'C', 'C', 'A', 'C', 'C', 'A', 'B', 'C', 'A', 'C', 'C', 'B', 'C', 'C']
---> 10 assert evolvedTest == targetArray
     11 # checks the initial array is left unchanged
     12 assert initialTestArray == ['A','B','C']*5

AssertionError:
```

Define a function that evolves a system that starts with initial amounts `NA`, `NB` and `NC` of $A$, $B$ and $C$ nuclei and evolved it in `n_timestep` from time $t=0$ to $t=t_{max}$. The function should return three arrays, one for each atom type, of the number of nuclei of that type at each time step. Each array should contain `n_timestep+1` elements including the initial amount. 
 

<!-- grade_id: evolve_system, points: 0, grade: false, solution: true -->
```python


def evolve_system(NA, NB, NC, rules, n_step):
    state = (['A'] * NA)+(['B'] * NB)+(['C'] * NC)

    A_count = numpy.zeros(n_step + 1, dtype=int)
    B_count = numpy.zeros(n_step + 1, dtype=int)
    C_count = numpy.zeros(n_step + 1, dtype=int)
    
    #print(A_count)
    #print(B_count)
    #print(C_count)
    
    
    A_count[0]=state.count('A')
    B_count[0]=state.count('B')
    C_count[0]=state.count('C')
    
    for i in range(1,n_step+1):
        state=evolveMany(state,rules)
        A_count[i]=state.count('A')
        B_count[i]=state.count('B')
        C_count[i]=state.count('C')
        #print(A_count)
        #print(B_count)
        #print(C_count)
        #print('\n')
        
    return A_count, B_count, C_count
```

<!-- grade_id: evolve_system_length, points: 2, grade: true, solution: false -->
```python
# these tests are worth 2 marks
rules = [
    ('A', 'B', 0.0033),
    ('B', 'C', 0.0075),
    ('C', 'A', 0.009)
    
]

r1, r2, r3 = evolve_system(0,0,250, rules, 17)
assert len(r1) == 18
assert len(r2) == 18
assert len(r3) == 18
```

<!-- grade_id: evolve_system_values, points: 2, grade: true, solution: false -->
```python
# these tests are worth 2 marks 
testrules = [
    ('A', 'B', 0.086),
    ('B', 'C', 0.075),
    ('C', 'A', 0.19)
    
]

random.seed(9485)
r1, r2, r3 = evolve_system(200,200,200, testrules, 20)
assert (r1 == [200, 213, 233, 250, 258, 251, 266, 263, 259, 260, 265, 259, 256,
        255, 258, 256, 259, 253, 249, 247, 253]).all()
assert (r2 == [200, 198, 201, 206, 205, 214, 214, 212, 216, 221, 225, 234, 236,
        238, 234, 235, 231, 245, 253, 256, 252]).all()
assert (r3 == [200, 189, 166, 144, 137, 135, 120, 125, 125, 119, 110, 107, 108,
        107, 108, 109, 110, 102,  98,  97,  95]).all()
```

## Plotting tasks

Create a plot with the number of $A$, $B$ and $C$ nuclei, starting with 250 $C$ nuclei and evolving the system for 100 hours using 200 steps and with neutron flux on. Evolve the system for another 100 hours (using 200 steps) without neutron flux (i.e. no transitions from $C$ to $A$).  

The half life of the $A$ atoms is 10.1 hours, the half life of $B$ nuclei is 15.7 hours and we can caracterise the rate of activation of $C$ into $A$ when the neutron flux is on with and effective half-life of 3.2 hours.

The plot should have the appropriate labels and legend. [8 marks]    




```python
nsteps = 200
t_total = 100
t_half_A = 10.1
t_half_B = 15.7
t_half_C = 3.2

```

Run the above simulation 20 times with 200 steps and use the results to calculate an average and the uncertainty on the number of $A$ atoms as a function of time. Use and `errorbar` plot for it. You might be interested in the `numpy.average` and `numpy.std` functions. The plot should have axis labels and a title.  [3 marks]


```python
nsim = 20

```